Heterogeneous efflux pump expression underpins phenotypic resistance to antimicrobial peptides

Author response:

The following is the authors’ response to the previous reviews.

Reviewer 1:

We would like to thank Reviewer 1 for recognising the importance of our findings on the heterogeneity in bacterial responses to tachyplesin.

(1) A double deletion of acrA and tolC (two out of the three components of the major constitutive RND efflux pump) reduces the appearance of the low accumulator phenotype, but interestingly, the single deletions have no effect, and a well-characterised inhibitor of RND efflux pumps also has no effect. The authors identify a two-component system, qseCB, that appears necessary for the appearance of low accumulators, but this system has pleiotropic effects on many cellular systems, with only tenuous connections to efflux. The selected pharmacological agents that could prevent the appearance of low accumulators do not offer clear insight into the mechanism by which low accumulators arise, because they have diverse modes of action.

We have added that “QseBC, was previously inferred to mediate resistance to a tachyplesin analogue by upregulating efflux genes based on transcriptomic analysis and hyper susceptibility of ΔqseBΔqseC mutants[113]”. However, we have also acknowledged that “it is conceivable that the deletion of QseBC has pleiotropic effects on other cellular mechanisms involved in tachyplesin accumulation.” and that “it is also conceivable that sertraline prevented the formation of the low accumulator phenotype via efflux independent mechanisms”

These amendments are reported on lines 525-527, 532-534 and 539-541 of our revised manuscript.

(2) The transcriptomics data collected for low and high accumulator sub-populations are interesting, but in my opinion, the conclusions that can be drawn from these data remain overstated. It is not possible to make any claims about the total amount of "protein synthesis, energy production, and gene expression" on the basis of RNA-Seq data. The reads from each sample are normalised, so there is no information about the total amount of transcript. Many elements of total cellular activity are post-transcriptionally regulated, so it is impossible to assess from transcriptomics alone. Finally, the transcriptomic data are analysed in aggregated clusters of genes that are enriched for biological processes, for example: "Cluster 2 included processes involved in protein synthesis, energy production, and gene expression that were downregulated to a greater extent in low accumulators than high accumulators". However, this obscures the fact that these clusters include genes that are generally inhibitory of the process named, as well as genes that facilitate the process.

We have now acknowledged that “that our data do not take into account post-transcriptional modifications that represent a second control point to survive external stressors.”

These amendments are reported on lines 534-535 of our revised manuscript.

The raw transcript counts can be found in Figure 3 – Source Data, we had added these data in our previous manuscript as requested by this reviewer.

We would also like to clarify that we have analysed our transcriptomic data via both clustering (i.e. Figure 3) and direct comparison of genes of interest (Table S1) and transcription factors (i.e. genes that are generally inhibitory of the process named, as well as genes that facilitate the process, Figure S12).

Finally, we would like to point out that in our revised manuscript (both this and its previous version) we are stating “Cluster 2 included processes involved in protein synthesis, energy production, and gene expression that were downregulated to a greater extent in low accumulators than high accumulators”. We do not think this is an overstatement, we do not use these data to make conclusions on the total amount of "protein synthesis, energy production, and gene expression".

(3) The authors have added an experiment to attempt to assess overall metabolic activity in the low accumulator and high accumulator populations, which is a welcome addition. They apply the redox dye resazurin and observe lower resorufin (reduced form) fluorescence in the low accumulator population, which they take to indicate a lower respiration rate. This seems possible, however, an important caveat is that they have shown the low accumulator population to retain substantially lower amounts of multiple different fluorescent molecules (tachyplesin-NBD, propidium iodide, ethidium bromide) intracellularly compared to the high accumulator population. It seems possible that the low accumulator population is also capable of removing resazurin or resorufin from the intracellular space, regardless of metabolic rate. Indeed, it has previously been shown that efflux by RND efflux pumps influences resazurin reduction to resorufin in both P. aeruginosa and E. coli. By measuring only the retained redox dye using flow cytometry, the results may be confounded by the demonstrated ability of the low accumulator population to remove various fluorescent dyes. More work is needed to strongly support broad conclusions about the physiological states of the low and high accumulator populations. The phenomenon of the emergence of low accumulators, which are phenotypically tolerant to the antimicrobial peptide tachyplesin, is interesting and important even if there is still work to be done to understand the mechanism by which it occurs.

We have now clarified that these assays were performed in the presence of 50 μM CCCP and that “CCCP was included to minimise differences in efflux activity and preserve resorufin retention between low and high accumulators, though some variability in efflux may still persist.” We have now added this information on lines 401-406. This information was only present in the caption of Figure S16 of our previous version of this manuscript.

We agree with the reviewers that more work needs to be done to fully understand this new phenomenon and we had already acknowledged in our previous version of this manuscript that other mechanisms could play a role in this new phenomenon, see lines 489-517 of the current manuscript.

Reviewer 2:

We would like to thank the reviewer for recognising that all their previous comments have now been satisfactorily addressed.

(1) Some mechanistic questions regarding tachyplesin-accumulation and survival remain. One general shortcoming of the setup of the transcriptomics experiment is that the tachyplesin-NBD probe itself has antibiotic efficacy and induces phenotypes (and eventually cell death) in the ´high accumulator´ cells. As the authors state themselves, this makes it challenging to interpret whether any differences seen between the two groups are causative for the observed accumulation pattern of if they are a consequence of differential accumulation and downstream phenotypic effects.

We agree with the reviewer and we had explicitly acknowledged this possibility on lines 281-285 (of the previous and current version of this manuscript).

(2) The statement ´ Moreover, we found that the fluorescence of low accumulators decreased over time when bacteria were treated with 20 μg mL´ is, in my opinion, not supported by the data shown in Figure S4C. That figure shows that the abundance of ´low accumulator´ cells decreases over time. Following the rationale that protease K treatment may cleave surface associated/ extracellular tachyplesin-NDB, this should lead to a shift of ´low accumulator´ population to the left, indicating reduced fluorescence intensity per cell. This is not so case, but the population just disappears. However, after 120 min of treatment more cells appear in the ´high accumulator´ state. This result is somewhat puzzling.

We agree with the reviewer that our previous discussion of this data could have been misleading. We have now reworded this part of the text as following: “We found that the fluorescence of high accumulators did not decrease over time when tachyplesin-NBD was removed from the extracellular environment and bacteria were treated with 20 μg mL-1 (0.7 μM) proteinase K, a widely-occurring serine protease that can cleave the peptide bonds of AMPs [43–45] (Figure S4B and C). These data suggest that tachyplesin-NBD primarily accumulates intracellularly in high accumulators.”

It is conceivable that extended exposure to proteinase K (i.e. we see a decrease in the abundance of low accumulators after 90 min treatment with proteinase K) increased the permeability to tachyplesin-NBD of low accumulators allowing tachyplesin-NBD to move from either the extracellular space or the membrane to the cell interior. However, we do not have data to prove this point.

Therefore, we have now removed our claim that the data obtained using proteinase K suggest that tachyplesin-NBD accumulates primarily in the membranes of low accumulators. We believe that our two separate microscopy analyses provide more direct, stronger and less ambiguous evidence that tachyplesin-NBD accumulates primarily in the membranes of low accumulators.

(3) The authors used the metabolic dye resazurin to measure the metabolic activity of low vs. high accumulators. I am not entirely convinced that the lower fluorescence resorufin fluorescence in tachyplesin-NBD accumulators really indicates lower metabolic activity, since a cell's fluorescence levels would also be affected by the cellular uptake and efflux. It appears plausible that the lower resorufin-fluorescence may result from reduced accumulation/increased efflux in the ‘low-tachyplesin NBD´ population.

We have now clarified that these assays were performed in the presence of 50 μM CCCP and that “CCCP was included to minimise differences in efflux activity and preserve resorufin retention between low and high accumulators, though some variability in efflux may still persist.” We have now added this information on lines 401-406. This information was only present in the caption of Figure S16 of our previous version of this manuscript.

(4) P8 line 343. The text should refer to Figure. 13B, instead of 14B

We have now changed the text accordingly on line 337.

Reviewer 3:

We would like to thank the reviewer for recognising that we have done a very impressive job in taking care of their comments.

(1) Despite these advances, the contribution of efflux may require more direct evidence to further dissect whether efflux is necessary, sufficient, or contributory. The facts that the key low efflux mutant still retains a small fraction of survivors and that the inhibitors used may cause other physiological changes leading to higher efflux are still unaccounted for. The lipidomic and vesicle findings, while intriguing, remain descriptive, and direct tests of their functional relevance would further solidify the mechanistic models.

We agree with the reviewers that more work needs to be done to fully understand this new phenomenon and we had already acknowledged in our previous version of this manuscript that other mechanisms could play a role in this new phenomenon, see lines 489-517 of the current manuscript.

Changes from the earlier version to the revised version

Abstract

Antimicrobial resistance threatens the viability of modern medical interventions. There is a dire need of developing novel approaches to counter resistance mechanisms employed by starved or slow-growing pathogens that are refractory to conventional antimicrobial therapies. Antimicrobial peptides have been advocated as potential therapeutic solutions due to low levels of genetic resistance observed in bacteria against these compounds. However, here we show that subpopulations of stationary phase Escherichia coli and Pseudomonas aeruginosa survive tachyplesin treatment without genetic mutations. These phenotypic variants display enhanced efflux activity to limit intracellular peptide accumulation. Differential regulation of genes involved in outer membrane vesicle section, membrane modification, and protease activity were also found between phenotypically resistant and susceptible cells. We discovered that formation of these phenotypic variants could be prevented by administering tachyplesin in combination with sertraline, a clinically used antidepressant, suggesting a novel approach for combatting antimicrobial-refractory stationary phase bacteria.

Introduction

Antimicrobial resistance (AMR) was estimated to claim 1.27 million lives in 2019 1 and is predicted to claim 10 million lives annually by 2050 2. Besides the deaths caused directly by infections, AMR also threatens the viability of many modern medical interventions such as surgery and organ transplant. This challenging situation has motivated the development of strategies to combat resistance mechanisms utilised by bacterial pathogens, particularly when they enter states of dormancy or slow growth and become more refractory to conventional antimicrobial treatments 3,4. In recent decades, natural defences such as bacteriocin proteins, peptides, endolysins, and antibodies have garnered substantial attention as potential clinical antimicrobial agents 5. Antimicrobial peptides (AMPs) have been advocated as potential therapeutic solutions to the AMR crisis. Notably, polymyxin-B and -E are in clinical use, and additional AMPs are undergoing clinical trials 6. AMPs can simultaneously target different cellular subsystems: they target the double membrane of Gram-negative bacteria by interacting with negatively charged membrane lipids while simultaneously disrupting cytoplasmic processes such as protein synthesis, DNA replication, cell wall biosynthesis, metabolism, and cell division 7.

Bacteria have developed genetic resistance to AMPs, including proteolysis by proteases, modifications in membrane charge and fluidity to reduce affinity, and extrusion by AMP transporters. However, compared to small molecule antimicrobials, AMP resistance genes typically confer smaller increases in resistance in experimental evolution analyses, with polymyxin-B and CAP18 being notable exceptions 8. Moreover, mobile resistance genes against AMPs are relatively rare and horizontal acquisition of AMP resistance is hindered by phylogenetic barriers owing to functional incompatibility with the new host bacteria 9. Plasmid-transmitted polymyxin resistance constitutes a notable exception 10, possibly because polymyxins are the only AMPs that have been in clinical use to date 9. However, whether bacterial populations harbour subpopulations that are transiently phenotypically resistant to AMPs without the acquisition of genetic mutations remains an open question. While there is limited understanding regarding AMPs and transient phenotypic resistance11, this phenomenon has been widely investigated in the case of small molecule antimicrobials and bacteriophages 12–17. Phenotypic resistance to antimicrobials is known to result in negative clinical outcomes and contributes to the emergence of genetic resistance 18–21. Therefore, further investigation into understanding the initial trajectories of the emergence of AMP resistance is required for the development of affordable, safe, and effective AMP treatments while simultaneously circumventing the pitfalls that have fuelled the current AMR crisis 22,23.

The 17 amino acid cationic β-hairpin AMP tachyplesin (2.27 kDa) 24 is a promising candidate AMP because it displays broad-spectrum, potent antibacterial and antifungal efficacy, minimal haemolytic effect, and minimal evolution of bacterial resistance 8. It has been speculated that tachyplesin simultaneously targets multiple cellular components 25,26. However, further investigation is needed to fully understand this aspect, along with exploring bacteria’s potential to transiently survive tachyplesin exposure without genetic mutations. Understanding the emergence of phenotypic resistance to tachyplesin could provide valuable insights into optimising the therapeutic potential of tachyplesin and other AMPs against microbial threats.

Achieving inhibitory concentrations proximal to its cellular target is crucial for tachyplesin (and antimicrobials in general) efficacy 27. At low concentrations, AMPs induce the formation of transient pores in bacterial membranes that allow transmembrane conduction of ions but not leakage of intracellular molecules 6,7. Beyond a critical ratio between the AMP and membrane lipid concentrations, these pores become stable, leading to leakage of cellular content, loss of transmembrane potential, and eventual cell death 6,7. Moreover, it is conceivable that tachyplesin, and other AMPs, must penetrate the bacterial inner membrane and accumulate at intracellular concentrations sufficient to interfere with cellular targets 28. Simultaneously, tachyplesin must avoid efflux via protein pumps, such as the Escherichia coli and Yersinia pestis tripartite pumps AcrAB-TolC and EmrAB-TolC, that confer genetic resistance against AMPs like protamine and polymyxin-B 29,30.

Crucially, recent discoveries have unveiled an association between reduced antimicrobial accumulation and transient phenotypic resistance to antimicrobials 31–35. Therefore, we hypothesise that phenotypic variations in the mechanisms regulating membrane lipid and protein compositions allow bacterial subpopulations to shift the balance between influx and efflux, reduce the intracellular accumulation of tachyplesin and survive treatment without the emergence of genetic mutations.

To test this hypothesis, we investigated the accumulation and efficacy of tachyplesin-1 against individual E. coli, Pseudomonas aeruginosa, Klebsiella pneumoniae and Staphylococcus aureus cells in their stationary phase of growth. To determine the mechanisms underpinning phenotypic resistance to tachyplesin, we performed comparative gene expression profiling and membrane lipid abundance analysis between subpopulations susceptible and transiently resistant to tachyplesin treatment. We used this knowledge to design a novel drug combination that increases tachyplesin efficacy in killing non-growing bacteria that are notoriously more refractory to treatment with AMPs or other antimicrobials. Our data demonstrate that streamlining antibacterial drug discovery could be facilitated by addressing the issue of low intracellular drug accumulation. We propose new combination treatments as a promising foundation for developing peptide-based drug combinations against dormant or slow growing bacteria.

Results

Tachyplesin accumulates heterogeneously within clonal E. coli and P. aeruginosa populations

To investigate the accumulation of tachyplesin in individual bacteria, we utilised a fluorescent derivative of the antimicrobial peptide, tachyplesin-1-nitrobenzoxadiazole (tachyplesin-NBD) 36 in combination with a flow cytometry assay adapted from previously reported protocols 37,38 (Figure S1A-C and Methods).

As expected, we found that bacterial fluorescence, due to accumulation of tachyplesin-NBD in individual stationary phase E. coli BW25113 cells, increased an order of magnitude when the extracellular concentration of tachyplesin-NBD increased from 0 to 8 μg mL-1 (0 to 3.2 μM). At the latter concentration, the single-cell fluorescence distribution displayed a median value of 2,400 a.u. after 60 min treatment at 37 °C (Figure 1A). However, a further increase in the extracellular concentration of tachyplesin-NBD to 16 μg mL-1 (6.3 μM) revealed a noticeable bimodal distribution in single-cell fluorescence. We classified bacteria in the distribution with a median value of 2,400 a.u. as low accumulators and those in the distribution with a median value of 28,000 a.u. as high accumulators (blue and red shaded areas, respectively, in Figure 1A and S1D-J).

Tachyplesin accumulates heterogeneously within clonal E. coli and P. aeruginosa populations.

(A) Tachyplesin-NBD accumulation in stationary phase E. coli BW25113 treated with increasing concentrations of tachyplesin-NBD for 60 min (0, 8, 16, 32, 46, 64, 128, and 256 μg mL-1 or 0.0, 3.2, 6.3, 12.7, 18.2, 25.4, 50.7, and 101.5 μM). Shaded regions show low (blue) and high (red) tachyplesin-NBD accumulation within an isogenic E. coli BW25113 population. Each histogram reports 10,000 recorded events and is a representative example of accumulation data from three independent biological replicates (Figure S1D-J). (B) Tachyplesin-NBD accumulation in nine stationary phase E. coli clinical isolates treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min. Accumulation data for E. coli BW25113 is reproduced from panel A. (C) Heatmap showing the presence of AMR genes and phylogenetic data of the clinical isolates. Corresponding gene products and the antimicrobials these genes confer resistance to are reported in Data Set S1. (D-F) Tachyplesin-NBD accumulation in P. aeruginosa (D), K. pneumoniae (E), and S. aureus (F) treated with 0 or 46 μg mL-1 (0 or 18.2 μM) tachyplesin-NBD for 60 min. Each histogram in each panel shows a representative example of accumulation data from three independent biological replicates. (G) Representative fluorescent images depicting low and high tachyplesin-NBD accumulators within an E. coli BW25113 population. Bacteria were continuously exposed to 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min in a microfluidic mother machine device. Blue stars and red diamonds indicate low and high accumulators, respectively. Scale bar: 3 μm. (H and I) Representative fluorescent images of individual low (H) and high (I) tachyplesin-NBD accumulators, respectively. Blue and red lines show a 2.2 μm-long cross-sectional line used for measuring fluorescence profile values in J and K, with the origin on the left side. (J and K) Normalised median (solid line), lower and upper quartiles (dotted lines) of fluorescence profile values of E. coli BW25113 cells plotted against the distance along the left side origin of a 2.2 μm-long straight line in low (J) and high (K) tachyplesin-NBD accumulators, respectively. Phenotype assignment was further verified via propidium iodide staining (see Figure 2G). Inset: each dot represents the corresponding membrane to cell centre fluorescence ratio measured at the points indicated in H and I, dashed lines indicate the median and quartiles of each distribution. Statistical significance was assessed using an unpaired nonparametric Mann-Whitney U test with a two-tailed p-value and confidence level at 95%. ****: p < 0.0001. Data was collected from three independent biological replicates.

A further increase in the extracellular concentration of tachyplesin-NBD to 46 μg mL-1 (18.2 μM) still yielded a bimodal distribution in single-cell fluorescence, 43% of the population being low accumulators; in contrast, only 5% of the population were low accumulators when tachyplesin-NBD treatment at 46 μg mL-1 (18.2 μM) was carried out against exponential phase E. coli (Figure S2). Notably, the minimum inhibitory concentration (MIC) of tachyplesin-NBD against exponential phase E. coli is 1 μg mL-1 (0.4 μM) 35. Therefore, our data suggests that even at extracellular concentrations surpassing growth-inhibitory levels, stationary phase E. coli populations harbour two distinct phenotypes with differing physiological states that respond to treatment differently with high accumulators displaying over 10-fold greater fluorescence than low accumulators.

Subsequently, we tested whether this newly observed bimodal distribution of tachyplesin-NBD accumulation was unique to the E. coli BW25113 strain. We found that the bimodal distribution in single-cell fluorescence was also present in nine E. coli clinical isolates (treated with tachyplesin-NBD at 46 μg mL-1 or 18.2 μM for 60 min) characterised by a diverse genetic background and various AMR genes, such as the polymyxin-E resistance gene mcr1 39 (Figure 1B and C, and Data Set S1). The median fluorescence of low accumulators in strains JS1147 and MC04960 were shifted to the left and to the right, respectively, compared to BW25113, whereas the median fluorescence of high accumulators for strain JS1060 was shifted to the left compared to BW25113. Remarkably, all isolates tested harboured low accumulators of tachyplesin-NBD.

Furthermore, to examine whether the bimodal distribution of tachyplesin-NBD accumulation also occurred in other highly virulent bacterial pathogens, we treated three members of the ESKAPE pathogens with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min while in their stationary phase of growth. In P. aeruginosa populations, we observed a bimodal distribution of tachyplesin-NBD accumulation with the presence of low accumulators (Figure 1D). In contrast, we did not find evidence of low tachyplesin-NBD accumulators in K. pneumoniae (Figure 1E) and recorded only low levels of tachyplesin-NBD accumulation in S. aureus (Figure 1F) in accordance with previous reports 40. Finally, we tested whether this newly observed bimodal distribution was a feature unique to tachyplesin or is widespread across different AMPs. We found that E. coli BW25113 displayed only low accumulators of another β-hairpin AMP, arenicin-NBD, and only high accumulators of cyclic lipopeptide AMPs, polymyxin-B-NBD and octapeptin-NBD (Figure S3).

These data demonstrate that two key pathogens display bimodal accumulation of tachyplesin, although only limited evolution of genetic resistance against tachyplesin has been previously observed 8, but not towards polymyxin-B against which they develop genetic resistance 8.

Tachyplesin accumulates primarily on the membranes of low accumulators

Earlier version: Next, we set out to characterise the differences between low and high tachyplesin accumulators. Firstly, we found that E. coli low accumulators treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min at 37 °C exhibited a fluorescence distribution akin to the entire population of cells treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD at 0 °C (Figure S4Α). At this low temperature, antimicrobials adhere non-specifically to bacterial surfaces as the passive and active transport across bacterial membranes is significantly diminished 38. Furthermore, post-binding peptide-lipid interactions are significantly diminished at 0 °C, but not the primary binding step 41,42. Moreover, we found that the fluorescence of low accumulators decreased over time when bacteria were treated with 20 μg mL-1 (0.7 μM) proteinase K, a widely-occurring serine protease that can cleave the peptide bonds of AMPs43–45; in contrast, the fluorescence of high accumulators did not decrease in the presence of proteinase K (Figure S4B and C). These data therefore suggest that tachyplesin-NBD is present only on the cell membranes of low accumulators and both on the membranes and intracellularly in high accumulators.

Revised version: Next, we set out to characterise the differences between low and high tachyplesin accumulators. Firstly, we found that E. coli low accumulators treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min at 37 °C exhibited a fluorescence distribution akin to the entire population of cells treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD at 0 °C (Figure S4Α). At this low temperature, antimicrobials adhere non-specifically to bacterial surfaces as the passive and active transport across bacterial membranes is significantly diminished 38. Furthermore, post-binding peptide-lipid interactions are significantly diminished at 0 °C, but not the primary binding step 41,42. Moreover, we found that the fluorescence of high accumulators did not decrease over time when tachyplesin-NBD was removed from the extracellular environment and bacteria were treated with 20 μg mL-1 (0.7 μM) proteinase K, a widely-occurring serine protease that can cleave the peptide bonds of AMPs 43–45 (Figure S4B and C). These data suggest that tachyplesin-NBD primarily accumulates intracellularly in high accumulators.

Earlier version: Next, we measured tachyplesin-NBD accumulation in individual E. coli cells using our microfluidics-based microscopy platform 46–48 with the well-established mother machine device 49,50. The mother machine is equipped with thousands of microfluidic channels that trap and host individual E. coli cells, enabling controlled medium exchange via pressure-driven microfluidics. Consistent with our flow cytometry data, we observed a bimodal accumulation of tachyplesin-NBD (Figure S5), with the presence of low and high accumulators (Figure 1G-I, blue stars and red diamonds, respectively). However, the relative abundance of low accumulators was lower in the microfluidics experiments possibly due to lower cell densities. Low accumulators exhibited a significantly higher membrane to cell centre fluorescence ratio compared to high accumulators (p-value < 0.0001, Figure 1J and 1K) and confocal microscopy further confirmed that tachyplesin-NBD is present only on the cell membranes of low accumulators and both on the membranes and intracellularly in high accumulators (Figure S6).

Revised version: Next, we measured tachyplesin-NBD accumulation in individual E. coli cells using our microfluidics-based microscopy platform 46–48 with the well-established mother machine device 49,50. The mother machine is equipped with thousands of microfluidic channels that trap and host individual E. coli cells, enabling controlled medium exchange via pressure-driven microfluidics. Consistent with our flow cytometry data, we observed a bimodal accumulation of tachyplesin-NBD (Figure S5), with the presence of low and high accumulators (Figure 1G-I, blue stars and red diamonds, respectively). However, the relative abundance of low accumulators was lower in the microfluidics experiments possibly due to lower cell densities. Low accumulators exhibited a significantly higher membrane to cell centre fluorescence ratio compared to high accumulators (p-value < 0.0001, Figure 1J and 1K) and confocal microscopy further confirmed that tachyplesin-NBD is present only on the cell membranes of low accumulators and primarily intracellularly in high accumulators (Figure S6).

We also compared the forward scatter and violet side scatter values, and cell lengths between low and high tachyplesin-NBD accumulators. We found that these two phenotypes had similar forward and violet side scatter values (Figure S7A and S7B), and high accumulators were significantly smaller than low accumulators in the microfluidics platform (p-value < 0.0001, Figure S7C). These data clearly exclude the possibility that high tachyplesin-NBD fluorescence is attributed to larger cell sizes.

Tachyplesin accumulation on the bacterial membranes is insufficient for bacterial eradication

To test the hypothesis that the accumulation of tachyplesin on the bacterial membranes may be insufficient for bacterial eradication, we conducted measurements of tachyplesin-NBD accumulation and efficacy in bacterial eradication.

Stationary phase E. coli was treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min followed by fluorescence-activated cell sorting to separate low and high accumulators, alongside untreated control cells (Figure 2A, 2B, S8, and Methods). Subsequently, we assessed the survival of each sorted sample via colony-forming unit assays and found that the survival fraction of low accumulators was not significantly different of that measured for untreated cells, whereas the survival fraction of high accumulators was significantly lower (p-value < 0.0001, Figure 2C).

Intracellular tachyplesin accumulation is essential for its antimicrobial efficacy.

(A and B) Fluorescence and side scatter values of individual E. coli treated with M9 at 37 °C for 60 min (A) and 46 μg mL-1 (18.2 μM) tachyplesin-NBD in M9 at 37 °C for 60 min (B). The black rectangles show the gates used to sort approximately one million untreated control cells, low and high tachyplesin-NBD accumulators for subsequent analysis. Data collected from four independent biological replicates. (C) Survival fraction of cells sorted through the untreated control, low and high tachyplesin-NBD accumulator gates presented in A and B. Bars and error bars represent the mean and standard deviation of data obtained from four biological replicates, each comprising three technical replicates. Statistical significance was assessed using an ordinary one-way ANOVA with Tukey’s multiple comparisons test and confidence level at 95%. ****: p < 0.0001, ns: p > 0.05. (D-F) Representative microscopic images depicting brightfield (D), tachyplesin-NBD fluorescence (E), and propidium iodide (PI) fluorescence (F) after exposure to 46 μg mL-1 (18.2 μM) tachyplesin-NBD in M9 for 60 min, followed by a M9 wash for 60 min, then 30 µM PI staining for 15 min at 37 °C in the microfluidic mother machine. Blue stars and red diamonds indicate low and high accumulators, respectively. Scale bar indicates 2.5 μm. The left part of Figure 2E is reproduced from part of Figure 1G. (G) Correlation between tachyplesin-NBD and PI fluorescence of N = 371 individual E. coli cells collected from three independent biological replicates. The purple dashed line shows a nonlinear regression (semi-log) of the data (r2 = 0.70). (H) Correlation between the proportion of low accumulators as a percentage of the whole bacterial population measured via flow cytometry and survival fraction measured via colony-forming unit assays. Symbols and error bars represent the mean and standard deviation of three independent biological replicates, each comprising three technical replicates. Tachyplesin-NBD treatment concentration indicated by colour gradient. Some error bars are masked by the data points. The purple dashed line illustrates a linear regression of the data (r2 = 0.98, p-value < 0.0001).

To investigate the hypothesis that tachyplesin causes less membrane damage to low accumulators, we employed our microfluidics-based microscopy platform to simultaneously measure tachyplesin-NBD accumulation and bacterial membrane integrity via propidium iodide (PI) 51. We found a strong positive semi-log correlation between tachyplesin-NBD fluorescence and PI fluorescence (r2 = 0.70, p-value < 0.0001): all bacteria exhibiting PI staining had a tachyplesin-NBD fluorescence greater than 800 a.u. and were high accumulators characterised by tachyplesin-NBD localisation both on the bacterial membranes and intracellularly; by contrast, low accumulators, where tachyplesin-NBD is primarily localised on the bacterial membranes, did not stain with PI (Figure 2D-G).

Finally, we set out to determine the contribution of low accumulators to the overall population survival. We exposed stationary phase E. coli to increasing concentrations of tachyplesin-NBD for 60 min, while measuring tachyplesin-NBD accumulation in individual bacteria and the bacterial eradication efficacy of tachyplesin-NBD at the population level. We used the accumulation data to measure the proportion of low accumulators relative to the whole population of bacteria at increasing tachyplesin-NBD concentrations and found a strong positive linear correlation with the population survival fraction (r2 = 0.98, p-value < 0.0001, Figure 2H). Taken together with PI staining data indicating membrane damage caused by high tachyplesin accumulation, these data demonstrate that low accumulators, which primarily retain tachyplesin-NBD on the bacterial membranes, maintain membrane integrity and strongly contribute to the survival of the bacterial population in response to tachyplesin treatment.

Enhanced efflux, membrane lipid alterations and increased secretion of outer membrane vesicles associated with low accumulation of tachyplesin

Next, we set out to investigate the mechanisms underpinning low accumulation of tachyplesin. Following our established protocols 52,53, we performed genome-wide comparative transcriptome analysis between low accumulators, a 1:1 (v/v) mixture of low and high accumulators, high accumulators, and untreated stationary phase bacteria that were sorted via fluorescence-activated cell sorting. Principal component analysis revealed clustering of the replicate transcriptomes of low tachyplesin-NBD accumulators, the 1:1 mixture of low and high accumulators, high accumulators, and untreated control cells with a distinct separation between these groups on the PC1 plane (Figure S9).

Earlier version: Our analysis revealed significant differential expression of 1,623 genes in at least one of the three groups compared to untreated cells (Data Set S2). By performing cluster analysis of these genes 48,53 (see Methods), we identified five distinct clusters of genes based on their expression profile across the three groups relative to the control population (Figure 3A, 3B, and S10). Gene ontology enrichment revealed significant enrichment of biological processes in clusters 2, 4, and 5 (Figure 3C and 3D). Genes in cluster 2 and 5 were downregulated to a greater extent in low accumulators; genes in cluster 4 were upregulated to a greater extent in low accumulators. This analysis enabled the identification of major transcriptional differences between low and high accumulators of tachyplesin.

Revised version: Our analysis revealed significant differential expression of 1,563 genes in at least one of the three groups compared to untreated cells (Data Set S2). By performing cluster analysis of these genes 48,53 (see Methods), we identified five distinct clusters of genes based on their expression profile across the three groups relative to the control population (Figure 3A, 3B, and S10). Gene ontology enrichment revealed significant enrichment of biological processes in clusters 2, 4, and 5 (Figure 3C and 3D). Genes in cluster 2 and 5 were downregulated to a greater extent in low accumulators; genes in cluster 4 were upregulated to a greater extent in low accumulators. This analysis enabled the identification of major transcriptional differences between low and high accumulators of tachyplesin.

Biological processes associated with low tachyplesin accumulation.

Earlier version: (A and B) Log2 fold changes in transcript reads of genes in low and high tachyplesin-NBD accumulators (blue and red violin plots, respectively) relative to untreated stationary phase E. coli bacteria in cluster 2 (A) and cluster 4 (B). Each dot represents a single gene, dashed lines indicate the median and quartiles of each distribution. Dotted lines represent a log2 fold change of 0. Statistical significance was tested using a paired two-tailed Wilcoxon nonparametric test (due to non-normally distributed data) with a two-tailed p-value and confidence level at 95%. ****: p-value < 0.0001. The full list of genes belonging to each cluster are reported in Data Set S2 and the violin plots of log2 fold changes for all clusters are shown in Figure S10. Data was collected from four independent biological replicates then pooled for analysis. (C and D) Corresponding gene ontology enrichment plots of biological processes enriched in cluster 2 (C) and 4 (D). Each dot represents a biological process with its size indicating the number of genes associated with each process and the colour indicating the false discovery rate (FDR). The lines and their thickness represent the abundance of mutual genes present within the connected biological processes. The enrichment plot for cluster 5 is not shown as only one biological process, “response to biotic stimulus”, was enriched. Full details about each process are reported in Data Set S3. (E) Relative abundance of lipid classes (PG, phosphatidylglycerol; FA, fatty acids; PE, phosphatidylethanolamine; LPE, lysophosphatidylethanolamine) in untreated bacteria and bacteria treated with 46 µg mL-1 (18.2 μM) tachyplesin-NBD for 60 min.

Revised version: (A and B) Log2 fold changes in transcript reads of genes in low and high tachyplesin-NBD accumulators (blue and red violin plots, respectively) relative to untreated stationary phase E. coli bacteria in cluster 2 (A) and cluster 4 (B). Each dot represents a single gene, dashed lines indicate the median and quartiles of each distribution. Dotted lines represent a log2 fold change of 0. Statistical significance was tested using a paired two-tailed Wilcoxon nonparametric test (due to non-normally distributed data) with a two-tailed p-value and confidence level at 95%. ****: p-value < 0.0001. The full list of genes belonging to each cluster are reported in Data Set S2and the violin plots of log2 fold changes for all clusters are shown in Figure S10. Data was collected from four independent biological replicates then pooled for analysis. (C and D) Corresponding gene ontology enrichment plots of biological processes enriched in cluster 2 (C) and 4 (D). Each dot represents a biological process with its size indicating the number of genes associated with each process and the colour indicating the false discovery rate (FDR). The lines and their thickness represent the abundance of mutual genes present within the connected biological processes. The enrichment plot for cluster 5 is not shown as only one biological process, “response to biotic stimulus”, was enriched. Full details about each process are reported in Data Set S3. (E) Relative abundance of lipid classes (PG, phosphatidylglycerol; FA, fatty acids; PE, phosphatidylethanolamine; LPE, lysophosphatidylethanolamine) in untreated bacteria and bacteria treated with 46 µg mL-1 (18.2 μM) tachyplesin-NBD for 60 min.

Cluster 2 included processes involved in protein synthesis, energy production, and gene expression that were downregulated to a greater extent in low accumulators than high accumulators (see Table S1 and Data Set S3 for a short and complete list of genes involved in these processes, respectively).

Cluster 4 included processes involved in transport of organic substances that were upregulated to a greater extent in low accumulators than high accumulators (Figure 3B and 3D). Notably, genes encoding major facilitator superfamily (MFS) and resistance-nodulation-division (RND) efflux pump components such as EmrK, MdtM and MdtEF, were upregulated in low accumulators (Table S1 and Data Set S3).

Low and high accumulators displayed differential expression of biological pathways facilitating the synthesis and assembly of the lipopolysaccharides (LPS), which is the first site of interaction between gram-negative bacteria and AMPs 54. For example, cardiolipin synthase B ClsB was contained in cluster 2 and downregulated in low accumulators (Table S1 and Data Set S3), suggesting that low accumulators might display a lower content of cardiolipin that could bind tachyplesin due to its negative charge 40,55. Moreover, we performed a comparative lipidomic analysis of untreated versus tachyplesin-NBD treated stationary phase E. coli using ultra-high performance liquid chromatography-quadrupole time-of-flight mass spectrometry 56. In accordance with our transcriptomic data, we found that treatment with tachyplesin-NBD, caused a strong decrease of phosphatidylglycerol (PG from 37% to 3%, i.e. the main constituent of cardiolipin), along with the emergence of lysophosphatidylethanolamines (LPE, 10%) and slight increases in PE and FA (Figure 3E, S11, and Table S2).

OmpA, OmpC, DegP, TolA, TolB, and Pal were downregulated in low accumulators, and their deletion has previously been associated with increased secretion of outer membrane vesicles (OMVs) that can confer resistance to AMPs and small molecule antibiotics 57,58. Moreover, NlpA, LysS, WaaCF, and Hns were upregulated in low accumulators, and their overexpression has been previously associated with enhanced OMV secretion 59,60 (Table S1 and Data Set S3).

Finally, low accumulators displayed an upregulation of peptidases and proteases compared to high accumulators, suggesting a potential mechanism for degrading tachyplesin (Table S1 and Data Set S3).

Noteworthy, tachyplesin-NBD has antibiotic efficacy (see Figure 2) and has an impact on the E. coli transcriptome (Figure 3). Therefore, we cannot conclude whether the transcriptomic differences reported above between low and high accumulators of tachyplesin-NBD are causative for the distinct accumulation patterns or if they are a consequence of differential accumulation and downstream phenotypic effects.

Next, we sought to infer transcription factor (TF) activities via differential expression of their known regulatory targets 61. A total of 126 TFs were inferred to exhibit differential activity between low and high accumulators (Data Set S4). Among the top ten TFs displaying higher inferred activity in low accumulators compared to high accumulators, four regulate transport systems, i.e. Nac, EvgA, Cra, and NtrC (Figure S12). However, further experiments should be carried out to directly measure the activity of these TFs.

Taken together these data suggest that phenotypic variants within clonal bacterial populations display a differential regulation of lipid composition, efflux, outer membrane vesicle secretion and proteolytic processes that affect tachyplesin accumulation levels.

Low accumulators display enhanced efflux activity

Next, we set out to delve deeper into the specific molecular mechanisms utilised by low accumulators in response to tachyplesin exposure. We investigated the accumulation kinetics of tachyplesin-NBD in individual stationary phase E. coli cells after different durations of tachyplesin-NBD treatment. Strikingly, within just 15 min of treatment, we observed high levels of tachyplesin-NBD accumulation in all cells, placing the whole population in the high accumulator group with a median fluorescence of 15,000 a.u. (Figure 4A). However, after 30 min of treatment, a subpopulation began to display lower fluorescence, and after 60 min, the low accumulator fluorescence distribution became evident with a median fluorescence of 3,600 a.u. (Figure 4A). By contrast, the median fluorescence of high accumulators increased from 15,000 a.u. at 15 min to 32,000 a.u. after 120 min (Figure 4A). Taken together, these data demonstrate that the entire isogenic E. coli population initially accumulates tachyplesin-NBD to high levels, but a subpopulation can reduce intracellular accumulation of the drug in response to treatment while the other subpopulation continues to accumulate the drug to higher levels. Reduced intracellular accumulation of tachyplesin-NBD in the presence of extracellular tachyplesin-NBD could be due to decreased drug influx, increased drug efflux, increased proteolytic activity, or increased secretion of OMVs.

Low accumulators of tachyplesin-NBD display enhanced efflux activity.

(A) Distribution of tachyplesin-NBD accumulation in stationary phase E. coli BW25113 over 240 min of treatment with 46 μg mL-1 (18.2 μM) tachyplesin-NBD in M9 at 37 °C. (B) Tachyplesin-NBD accumulation in E. coli BW25113 single or double gene-deletion mutants lacking efflux components or regulators of efflux components. Stationary phase populations of each mutant were treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD at 37 °C for 60 min and single-cell fluorescence was measured via flow cytometry. (C-D) Efflux of tachyplesin-NBD over 240 min, after an initial 15 min preloading with 46 μg mL-1 (18.2 μM) tachyplesin-NBD in M9 then washed and transferred into either M9 (C), or sertraline (30 μg mL-1 or 98 μM) (D). The top histograms of panels B-D are reproduced from the 10 min histogram of panel A. Histograms in panels A-D are representative of three independent biological replicates and report 10,000 events. (E and F) Correlation of tachyplesin-NBD and PI fluorescence of individual stationary phase E. coli BW25113 cells measured in the microfluidic mother machine after 60 min of 46 μg mL-1 (18.2 μM) tachyplesin-NBD treatment (N = 108) (E), or 46 μg mL-1 (18.2 μM) tachyplesin-NBD cotreatment with 30 μg mL-1 (98 μM) sertraline (N = 108) (F). Purple dashed lines show nonlinear (semi-log) regressions (r2 = 0.74 and 0.38, respectively). (G) Survival fraction of stationary phase E. coli BW25113 over 240 min treatment with tachyplesin-1 (46 μg mL-1 or 20.3 μM, magenta squares), sertraline (30 μg mL-1 or 98 μM, green circles), tachyplesin-1 (46 μg mL-1 or 20.3 μM) cotreatment with sertraline (30 μg mL-1 or 98 μM) (purple triangles), or incubation in M9 (black diamonds). Symbols and error bars indicate the mean and standard deviation of measurements performed in three independent biological replicates comprising three technical replicates each. Some error bars are masked behind the symbols.

Earlier version: Next, we performed efflux assays using ethidium bromide (EtBr) by adapting a previously described protocol 62. Briefly, we preloaded stationary phase E. coli with EtBr by incubating cells at a concentration of 254 µM EtBr in M9 medium for 90 min. Cells were then pelleted and resuspended in M9 to remove extracellular EtBr. Single-cell EtBr fluorescence was measured at regular time points in the absence of extracellular EtBr using flow cytometry. This analysis revealed a progressive homogeneous decrease of EtBr fluorescence due to efflux from all cells within the stationary phase E. coli population (Figure S13A). In contrast, when we performed efflux assays by preloading cells with tachyplesin-NBD (46 μg mL-1 or 18.2 μM), followed by pelleting and resuspension in M9 to remove extracellular tachyplesin-NBD, we observed a heterogeneous decrease in tachyplesin-NBD fluorescence in the absence of extracellular tachyplesin-NBD: a subpopulation retained high tachyplesin-NBD fluorescence, i.e. high accumulators; whereas another subpopulation displayed decreased tachyplesin-NBD fluorescence, 60 min after the removal of extracellular tachyplesin-NBD (Figure 4B). Since these assays were performed in the absence of extracellular tachyplesin-NBD, decreased tachyplesin-NBD fluorescence could not be ascribed to decreased drug influx or increased secretion of OMVs in low accumulators, but could be due to either enhanced efflux or proteolytic activity in low accumulators.

Revised version: Next, we performed efflux assays using ethidium bromide (EtBr) by adapting a previously described protocol 62. Briefly, we preloaded stationary phase E. coli with EtBr by incubating cells at a concentration of 254 µM EtBr in LB for 90 min. Cells were then pelleted and resuspended in M9 to remove extracellular EtBr. Single-cell EtBr fluorescence was measured at regular time points in the absence of extracellular EtBr using flow cytometry. This analysis revealed a progressive homogeneous decrease of EtBr fluorescence due to efflux from all cells within the stationary phase E. coli population (Figure S13A). In contrast, when we performed efflux assays by preloading cells with tachyplesin-NBD (46 μg mL-1 or 18.2 μM), followed by pelleting and resuspension in M9 to remove extracellular tachyplesin-NBD, we observed a heterogeneous decrease in tachyplesin-NBD fluorescence in the absence of extracellular tachyplesin-NBD: a subpopulation retained high tachyplesin-NBD fluorescence, i.e. high accumulators; whereas another subpopulation displayed decreased tachyplesin-NBD fluorescence, 60 min after the removal of extracellular tachyplesin-NBD (Figure 4B). Since these assays were performed in the absence of extracellular tachyplesin-NBD, decreased tachyplesin-NBD fluorescence could not be ascribed to decreased drug influx or increased secretion of OMVs in low accumulators, but could be due to either enhanced efflux or proteolytic activity in low accumulators.

Earlier version: Next, we repeated efflux assays using EtBr in the presence of 46 μg mL-1 (or 20.3 µM) extracellular tachyplesin-1. We observed a heterogeneous decrease of EtBr fluorescence with a subpopulation retaining high EtBr fluorescence (i.e. high tachyplesin accumulators) and another population displaying reduced EtBr fluorescence (i.e. low tachyplesin accumulators, Figure S14B) when extracellular tachyplesin-1 was present. Moreover, we repeated tachyplesin-NBD efflux assays in the presence of M9 containing 50 μg mL-1 (244 μM) carbonyl cyanide m-chlorophenyl hydrazone (CCCP), an ionophore that disrupts the proton motive force (PMF) and is commonly employed to abolish efflux and found that all cells retained tachyplesin-NBD fluorescence (Figure S15B). However, it is important to note that CCCP does not only abolish efflux but also other respiration-associated and energy-driven processes 63.

Revised version: Next, we repeated efflux assays using EtBr in the presence of 46 μg mL-1 (or 20.3 µM) extracellular tachyplesin-1. We observed a heterogeneous decrease of EtBr fluorescence with a subpopulation retaining high EtBr fluorescence (i.e. high tachyplesin accumulators) and another population displaying reduced EtBr fluorescence (i.e. low tachyplesin accumulators, Figure S13B) when extracellular tachyplesin-1 was present. Moreover, we repeated tachyplesin-NBD efflux assays in the presence of M9 containing 50 μg mL-1 (244 μM) carbonyl cyanide m-chlorophenyl hydrazone (CCCP), an ionophore that disrupts the proton motive force (PMF) and is commonly employed to abolish efflux and found that all cells retained tachyplesin-NBD fluorescence (Figure S15B). However, it is important to note that CCCP does not only abolish efflux but also other respiration-associated and energy-driven processes 63.

Taken together, our data demonstrate that in the absence of extracellular tachyplesin, stationary phase E. coli homogeneously efflux EtBr, whereas only low accumulators are capable of performing efflux of intracellular tachyplesin after initial tachyplesin accumulation. In the presence of extracellular tachyplesin, only low accumulators can perform efflux of both intracellular tachyplesin and intracellular EtBr. However, it is also conceivable that besides enhanced efflux, low accumulators employ proteolytic activity, OMV secretion, and variations to their bacterial membrane to hinder further uptake and intracellular accumulation of tachyplesin in the presence of extracellular tachyplesin.

Next, we performed tachyplesin-NBD accumulation assays using 28 single and 4 double E. coli BW25113 gene-deletion mutants of efflux components and transcription factors regulating efflux. While for the majority of the mutants we recorded bimodal distributions of tachyplesin-NBD accumulation similar to the distribution recorded for the E. coli BW25113 parental strain (Figure 4B and Figure S13), we found unimodal distributions of tachyplesin-NBD accumulation constituted only of high accumulators for both ΔqseB and ΔqseBΔqseC mutants as well as reduced numbers of low accumulators for the ΔacrAΔtolC mutant (Figure 4B). Considering that the AcrAB-TolC tripartite RND efflux system is known to confer genetic resistance against AMPs like protamine and polymyxin-B 29,30 and that the quorum sensing regulators qseBC might control the expression of acrA 64, these data further corroborate the hypothesis that low accumulators can efflux tachyplesin and survive treatment with this AMP.

Tachyplesin accumulation and efficacy can be boosted with the addition of efflux pump inhibitors or nutrients

Earlier version: Therefore, we screened a panel of seven alternative efflux pump inhibitors (EPIs) to identify compounds capable of preventing the formation of low accumulators of tachyplesin-NBD. Interestingly, M9 containing 30 µg mL-1 (98 μM) sertraline (Figure 4D and S15C), an antidepressant which inhibits efflux activity of RND pumps, potentially through direct binding to efflux pumps 65 and decreasing the PMF 66, or 50 µg mL-1 (110 μM) verapamil (Figure S15D), a calcium channel blocker that inhibits MATE transporters 67 by a generally accepted mechanism of PMF generation interference 68,69, was able to prevent the emergence of low accumulators. Furthermore, tachyplesin-NBD cotreatment with sertraline simultaneously increased tachyplesin-NBD accumulation and PI fluorescence levels in individual cells (Figure 4E and F, p-value < 0.0001 and 0.05, respectively). The use of berberine, a natural isoquinoline alkaloid that inhibits MFS transporters 70 and RND pumps 71, potentially by inhibiting conformational changes required for efflux activity 70, and baicalein, a natural flavonoid compound that inhibits ABC 72 and MFS 73,74 transporters, potentially through PMF dissipation 75, prevented the formation of a bimodal distribution of tachyplesin accumulation, however displayed reduction in fluorescence of the whole population (Figure S15E and F). Phenylalanine-arginine beta-naphthylamide (PAαN), a synthetic peptidomimetic compound that inhibits RND pumps 76 through competitive inhibition 77, reserpine, an indole alkaloid that inhibits ABC and MFS transporters, and RND pumps 78, by altering the generation of the PMF 69, and 1-(1-naphthylmethyl)piperazine (NMP), a synthetic piperazine derivative that inhibits RND pumps 79, through non-competitive inhibition 80, did not prevent the emergence of low accumulators (Figure S15G-I).

Revised version: Therefore, we screened a panel of seven alternative efflux pump inhibitors (EPIs) to identify compounds capable of preventing the formation of low accumulators of tachyplesin-NBD. Interestingly, M9 containing 30 µg mL-1 (98 μM) sertraline (Figure 4D and S15C), an antidepressant which inhibits efflux activity of RND pumps, potentially through direct binding to efflux pumps 65 and decreasing the PMF 66, or 50 µg mL-1 (110 μM) verapamil (Figure S15D), a calcium channel blocker that inhibits MATE transporters 67 by a generally accepted mechanism of PMF generation interference 68,69, was able to prevent the emergence of low accumulators. Furthermore, tachyplesin-NBD cotreatment with sertraline simultaneously increased tachyplesin-NBD accumulation levels in individual cells and proportion of PI stained cells (Figure 4E and F). The use of berberine, a natural isoquinoline alkaloid that inhibits MFS transporters 70 and RND pumps 71, potentially by inhibiting conformational changes required for efflux activity 70, and baicalein, a natural flavonoid compound that inhibits ABC 72 and MFS 73,74 transporters, potentially through PMF dissipation 75, prevented the formation of a bimodal distribution of tachyplesin accumulation, however displayed reduction in fluorescence of the whole population (Figure S15E and F). Phenylalanine-arginine beta-naphthylamide (PAβN), a synthetic peptidomimetic compound that inhibits RND pumps 76 through competitive inhibition 77, reserpine, an indole alkaloid that inhibits ABC and MFS transporters, and RND pumps 78, by altering the generation of the PMF 69, and 1-(1-naphthylmethyl)piperazine (NMP), a synthetic piperazine derivative that inhibits RND pumps 79, through non-competitive inhibition 80, did not prevent the emergence of low accumulators (Figure S15G-I).

Next, we investigated whether sertraline could enhance the bactericidal activity of the parent drug, tachyplesin-1. Using colony-forming unit assays, we measured the survival fraction of stationary phase E. coli treated with either tachyplesin-1 at 46 μg mL-1 (20.3 μM) (which has an MIC value of 1 µg mL-1 or 0.43 μM against E. coli 35) or sertraline at 30 μg mL-1 (98 μM) (for which we determined an MIC value of 128 μg mL-1 or 418 μM against E. coli, see Figure 4 – Source Data), or a combination of both compounds, and compared these data to untreated E. coli incubated in M9. The survival fraction after treatment with sertraline was comparable to that measured for untreated E. coli, whereas the survival fraction measured after the combination treatment was five-fold lower than that measured after tachyplesin-1 treatment (Figure 4G). Increased tachyplesin efficacy in the presence of sertraline is likely due to efflux inhibition. However, it is also conceivable that increased tachyplesin efficacy is due to metabolic and transcriptomic changes induced by sertraline.

Earlier version: Our transcriptomics analysis also showed that low tachyplesin accumulators downregulated protein synthesis, energy production, and gene expression compared to high accumulators. To gain further insight on the metabolic state of low tachyplesin accumulators, we employed the membrane-permeable redox-sensitive dye, resazurin, which is reduced to the highly fluorescent resorufin in metabolically active cells. We first treated stationary phase E. coli with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min, then washed the cells, and then incubated them in 1 μM resazurin for 15 min and measured single-cell fluorescence of resorufin and tachyplesin-NBD simultaneously via flow cytometry. We found that low tachyplesin-NBD accumulators also displayed low fluorescence of resorufin, whereas high tachyplesin-NBD accumulators also displayed high fluorescence of resorufin (Figure S16), suggesting lower metabolic activity in low tachyplesin-NBD accumulators.

Revised version: Our transcriptomics analysis also showed that low tachyplesin accumulators downregulated protein synthesis, energy production, and gene expression compared to high accumulators. To gain further insight on the metabolic state of low tachyplesin accumulators, we employed the membrane-permeable redox-sensitive dye, resazurin, which is reduced to the highly fluorescent resorufin in metabolically active cells. We first treated stationary phase E. coli with 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min, then washed the cells, and then incubated them in 1 μM resazurin and 50 μM CCCP for 15 min and measured single-cell fluorescence of resorufin and tachyplesin-NBD simultaneously via flow cytometry. CCCP was included to minimise differences in efflux activity and preserve resorufin retention between low and high accumulators, though some variability in efflux may still persist. We found that low tachyplesin-NBD accumulators also displayed low fluorescence of resorufin, whereas high tachyplesin-NBD accumulators also displayed high fluorescence of resorufin (Figure S16), suggesting lower metabolic activity in low tachyplesin-NBD accumulators.

Therefore, we set out to test the hypothesis that the formation of low accumulators could be prevented by manipulating the nutrient environment around the bacteria. We treated stationary phase E. coli with 46 μg mL-1 (18.2 μM) tachyplesin-NBD either in M9, or M9 supplemented with 0.4% glucose and 0.2% casamino acids. After 15 min of treatment, both conditions resulted in a unimodal fluorescence distribution, with median values of 15,000 a.u. and 18,000 a.u., respectively (Figure 4A and 5A, respectively). However, upon extending the treatment to 30 min, we observed distinct responses to tachyplesin-NBD treatment. For E. coli treated with tachyplesin-NBD in M9, the emergence of low accumulators led to a bimodal distribution of single-cell fluorescence (Figure 4A). In contrast, when E. coli was treated with tachyplesin-NBD in M9 supplemented with nutrients, the fluorescence of the entire population increased, resulting in a unimodal distribution of single-cell fluorescence with a median of 35,000 a.u. (Figure 5A).

Impact of nutritional environment and growth phase on tachyplesin-NBD accumulation and efficacy.

(A and B) Temporal dependence of the distribution of tachyplesin-NBD accumulation in stationary phase E. coli in M9 supplemented with 0.4% glucose and 0.2% casamino acids (A) and exponential phase E. coli in M9 only (B) treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD. Each distribution is representative of three independent biological replicates. (C) Survival fraction of untreated stationary phase E. coli in M9 only (black diamonds), and bacteria treated with 46 μg mL-1 (18.2 μM) tachyplesin-NBD over 240 min in stationary phase E. coli in M9 only (green circles), stationary phase E. coli treated in M9 supplemented with 0.4% glucose and 0.2% casamino acids (purple triangles), exponential phase E. coli treated in M9 only (magenta squares). Symbols and error bars indicate the mean and standard deviation of three biological replicates measured from three technical replicates each. Some error bars are masked behind the symbols.

Similarly, E. coli incubated in LB for 3 h before drug treatment, and therefore in the exponential phase of growth, displayed a unimodal fluorescence distribution with a median of 80,000 a.u. after 10 min treatment in tachyplesin-NBD in M9 (Figure 5B). These exponential phase bacteria also displayed a modest increase in cell size compared to stationary phase E. coli (Figure S17). Interestingly, a small subpopulation of low accumulators emerged within exponential phase E. coli after 120 min tachyplesin-NBD treatment in M9, with a median of 4,700 a.u. (Figure 5B). This subpopulation contributed to 10% of the entire isogenic E. coli population, five-fold less abundant compared to stationary phase E. coli treated in tachyplesin-NBD in M9 (Figure 4A). Finally, we investigated whether preventing the formation of low accumulators via nutrient supplementation could enhance the efficacy of tachyplesin treatment by performing colony-forming unit assays. We observed that tachyplesin exhibited enhanced efficacy in eradicating E. coli both when M9 was supplemented with nutrients, and against exponential phase E. coli compared to stationary phase E. coli without nutrient supplementation (Figure 5C).

Discussion

The link between antimicrobial accumulation and efficacy has not been widely investigated, especially in the context of AMPs 27. In fact, it has been widely accepted that the cell membrane is the primary target of most AMPs and that their main mechanism of action is membrane disruption. However, emerging evidence suggests that AMPs may also enter the intracellular environment and target intracellular processes, creating opportunities for the development of bacterial resistance 6,81–83.

Earlier version: Here, we demonstrate for the first time that there is a strong correlation between tachyplesin intracellular accumulation and efficacy. We show that after an initial homogeneous tachyplesin accumulation within a stationary phase E. coli population, tachyplesin is retained intracellularly by bacteria that do not survive tachyplesin exposure, whereas tachyplesin is retained only in the membrane of bacteria that survive tachyplesin exposure. These findings align with evidence showing that tachyplesin interacts with intracellular targets such as the minor groove of DNA 26,84, intracellular esterases 85, or the 3-ketoacyl carrier protein reductase FabG 25. However, our transcriptomic analysis did not reveal differential regulation of any of these pathways between low and high accumulators of tachyplesin, suggesting that bacteria likely employ alternative mechanisms rather than target downregulation to survive tachyplesin treatment, a point on which we expand below.

Revised version: Here, we demonstrate for the first time that there is a strong correlation between tachyplesin intracellular accumulation and efficacy. We show that after an initial homogeneous tachyplesin accumulation within a stationary phase E. coli population, tachyplesin is retained intracellularly by bacteria that do not survive tachyplesin exposure, whereas tachyplesin is retained only in the membrane of bacteria that survive tachyplesin exposure.

Added in the revised version: These findings align with evidence showing that modifications to tachyplesin that preserve membrane disruption but reduce membrane translocation, such as linear analogues or PEGylation, result in diminished antimicrobial activity 26,84. Furthermore, tachyplesin has also been reported to interact with intracellular targets such as the minor groove of DNA 26,85, intracellular esterases 86, or the 3-ketoacyl carrier protein reductase FabG 25. However, our transcriptomic analysis did not reveal differential regulation of any of these pathways between low and high accumulators of tachyplesin, suggesting that bacteria likely employ alternative mechanisms rather than target downregulation to survive tachyplesin treatment, a point on which we expand below.

Earlier version: The emergence of genetic resistance to AMPs is strikingly lower compared to resistance to small molecule antibiotics 8,83,86. However, research on AMP resistance has often focused on whole population responses 8,83,86,87, overlooking heterogeneous responses to drug treatment in isogenic bacteria. Crucially, phenotypic variants can survive treatment with small molecule antibiotics through several distinct phenomena, including persistence 88, the viable but nonculturable state 89, heteroresistance 90, tolerance 91, and perseverance 19.

Revised version: The emergence of genetic resistance to AMPs is strikingly lower compared to resistance to small molecule antibiotics 8,83,87. However, research on AMP resistance has often focused on whole population responses 8,83,87,88, overlooking heterogeneous responses to drug treatment in isogenic bacteria. Crucially, phenotypic variants can survive treatment with small molecule antibiotics through several distinct phenomena, including persistence 89, the viable but nonculturable state 90, heteroresistance 91, tolerance 92, and perseverance 19.

Here, we report the first observation of phenotypic resistance to tachyplesin within isogenic E. coli and P. aeruginosa populations. This novel phenomenon is characterised by a bimodal distribution of tachyplesin accumulation, revealing two distinct subpopulations: low (surviving) and high (susceptible) accumulators. This bimodal distribution draws similarities to the perseverance phenomenon identified for two antimicrobials possessing intracellular targets, rifampicin and nitrofurantoin. Perseverance is characterised by a bimodal distribution of single-cell growth rates, where the slower-growing subpopulation maintains growth during drug treatment 19.

Earlier version: Bacteria can reduce the intracellular concentration of antimicrobials via active efflux, allowing them to survive in the presence of high extracellular antimicrobial concentrations 92. Efflux as a mechanism of tachyplesin resistance has been shown at the population level in tachyplesin-resistant P. aeruginosa via transcriptome analysis 87,93, whereas phenotypic resistance to tachyplesin has not been investigated. Heterogeneous expression of efflux pumps within isogenic bacterial populations has been reported 31,34,35,50,94,95. However, recent reports have suggested that efflux is not the primary mechanism of antimicrobial resistance within stationary phase bacteria 33,96 and the impact of heterogeneous efflux pump expression on AMP accumulation is unknown. Here, we show for the first time the upregulation of a range of MFS, ABC, and RND efflux pumps in a stationary phase subpopulation that exhibited low accumulation of tachyplesin with corresponding enhanced survival and that the deletion of two components of the AcrAB-TolC efflux pump or the quorum sensing regulator QseBC leads to a dramatic decrease in the numbers of low accumulators.

Revised version: Bacteria can reduce the intracellular concentration of antimicrobials via active efflux, allowing them to survive in the presence of high extracellular antimicrobial concentrations 93. Efflux as a mechanism of tachyplesin resistance has been shown at the population level in tachyplesin-resistant P. aeruginosa via transcriptome analysis 88,94, whereas phenotypic resistance to tachyplesin has not been investigated. Heterogeneous expression of efflux pumps within isogenic bacterial populations has been reported 31,34,35,50,95,96. However, recent reports have suggested that efflux is not the primary mechanism of antimicrobial resistance within stationary phase bacteria 33,97 and the impact of heterogeneous efflux pump expression on AMP accumulation is unknown. Here, we show for the first time the upregulation of a range of MFS, ABC, and RND efflux pumps in a stationary phase subpopulation that exhibited low accumulation of tachyplesin with corresponding enhanced survival and that the deletion of two components of the AcrAB-TolC efflux pump or the quorum sensing regulator QseBC leads to a dramatic decrease in the numbers of low accumulators.

Earlier version: Moreover, whether certain resistance mechanisms utilised by phenotypic variants surviving antimicrobial treatment are preemptive or induced by drug treatment is underinvestigated 88,97 primarily due to the requirement of prior selection for phenotypic variants 31. As our accumulation assay did not require the prior selection for phenotypic variants, we have demonstrated that low accumulators emerge subsequent to the initial high accumulation of tachyplesin-NBD, suggesting enhanced efflux as an induced response. However, it is conceivable that other pre-existing traits of low accumulators also contribute to reduced tachyplesin accumulation. For example, reduced protein synthesis, energy production, and gene expression in low accumulators could slow down tachyplesin efficacy, giving low accumulators more time to mount efflux as an additional protective response.

Revised version: Moreover, whether certain resistance mechanisms utilised by phenotypic variants surviving antimicrobial treatment are preemptive or induced by drug treatment is underinvestigated 89,98 primarily due to the requirement of prior selection for phenotypic variants 31. As our accumulation assay did not require the prior selection for phenotypic variants, we have demonstrated that low accumulators emerge subsequent to the initial high accumulation of tachyplesin-NBD, suggesting enhanced efflux as an induced response. However, it is conceivable that other pre-existing traits of low accumulators also contribute to reduced tachyplesin accumulation. For example, reduced protein synthesis, energy production, and gene expression in low accumulators could slow down tachyplesin efficacy, giving low accumulators more time to mount efflux as an additional protective response.

Earlier version: Furthermore, most AMPs are cationic and interact electrostatically with the negatively charged LPS in the bacterial outer membrane 98. To reduce affinity to AMPs, bacteria can modify their membrane charge, thickness, fluidity, or curvature 99. At the population level, bacteria have been reported to add positively charged L-Ara4N or phosphoethanolamine (pEtN) to LPS lipid-A to resist AMPs 83. Accordingly, our transcriptomics analysis revealed the upregulation in the low accumulator subpopulation of the arn operon that is involved in the biosynthesis and attachment of the positively charged L-Ara4N to lipid-A 100. Furthermore, tachyplesin has been shown to bind strongly to negatively charged PG lipids 55. Accordingly, our lipidomics data revealed that tachyplesin-treated bacteria were composed of significantly lower abundance of PG lipids compared to the untreated bacteria.

Revised version: Furthermore, most AMPs are cationic and interact electrostatically with the negatively charged LPS in the bacterial outer membrane 99. To reduce affinity to AMPs, bacteria can modify their membrane charge, thickness, fluidity, or curvature 100. At the population level, bacteria have been reported to add positively charged L-Ara4N or phosphoethanolamine (pEtN) to LPS lipid-A to resist AMPs 83. Accordingly, our transcriptomics analysis revealed the upregulation in the low accumulator subpopulation of the arn operon that is involved in the biosynthesis and attachment of the positively charged L-Ara4N to lipid-A 101. Furthermore, tachyplesin has been shown to bind strongly to negatively charged PG lipids 55. Accordingly, our lipidomics data revealed that tachyplesin-treated bacteria were composed of significantly lower abundance of PG lipids compared to the untreated bacteria.

Earlier version: OMVs, nano-sized proteoliposomes secreted by Gram-negative bacteria, are a less studied system that allows bacteria to remove AMPs from the cell 101. Increased secretion of OMVs has also been shown to be induced by AMP exposure 58. OMVs are suggested to counteract AMPs via at least three mechanisms: the adsorption of extracellular AMPs to decrease AMP concentrations in the environment 57,58; removal of parts of the bacterial outer membrane affected by AMPs 58; and export of antimicrobials from the intracellular compartment 102. Our transcriptomic data revealed the downregulation in low accumulators of genes, such as ompAC, known to increase the secretion of OMVs upon deletion and the upregulation of genes, such as nlpA, whose overexpression has been linked to enhanced OMV secretion103,104.

Revised version: OMVs, nano-sized proteoliposomes secreted by Gram-negative bacteria, are a less studied system that allows bacteria to remove AMPs from the cell 102. Increased secretion of OMVs has also been shown to be induced by AMP exposure 58. OMVs are suggested to counteract AMPs via at least three mechanisms: the adsorption of extracellular AMPs to decrease AMP concentrations in the environment 57,58; removal of parts of the bacterial outer membrane affected by AMPs 58; and export of antimicrobials from the intracellular compartment 103. Our transcriptomic data revealed the downregulation in low accumulators of genes, such as ompAC, known to increase the secretion of OMVs upon deletion and the upregulation of genes, such as nlpA, whose overexpression has been linked to enhanced OMV secretion 104,105.

Earlier version: Bacteria also employ extracellular or intracellular proteases 7,83 and peptidases 105,106 to cleave and neutralise the activity of AMPs including tachyplesin 107. Our transcriptomic analysis revealed the simultaneous upregulation in low accumulators of proteases and peptidases that could potentially degrade tachyplesin either extracellularly or intracellularly.

Revised version: Bacteria also employ extracellular or intracellular proteases 7,83 and peptidases 106,107 to cleave and neutralise the activity of AMPs including tachyplesin 108. Our transcriptomic analysis revealed the simultaneous upregulation in low accumulators of proteases and peptidases that could potentially degrade tachyplesin either extracellularly or intracellularly.

Earlier version: Responses to extrinsic environmental cues, such as the starvation-induced stringent response, have been shown to induce failure in antimicrobial treatment 108. These responses are typically mediated by two-component regulatory systems (TCSs) sensing external stimuli via sensor kinases and altering gene expression via response regulators 109. Additionally, nucleoid associated proteins (NAPs) influence bacterial chromosome organisation and gene transcription in response to intracellular or extracellular stress factors 110. Finally, transcriptional regulators such as regulators of the LysR family and histone-like proteins (H-NS) also modulate gene expression 111. Interestingly the QseCB system has been demonstrated to facilitate resistance to a tachyplesin analogue 112. Accordingly, we demonstrated that mutants lacking qseB display only high accumulators of tachyplesin-NBD, suggesting that this two-component system might play an important role in the control of tachyplesin efflux, possibly by regulating the expression of AcrA or other efflux components64.

Revised version: Responses to extrinsic environmental cues, such as the starvation-induced stringent response, have been shown to induce failure in antimicrobial treatment 109. These responses are typically mediated by two-component regulatory systems (TCSs) sensing external stimuli via sensor kinases and altering gene expression via response regulators 110. Additionally, nucleoid associated proteins (NAPs) influence bacterial chromosome organisation and gene transcription in response to intracellular or extracellular stress factors 111. Finally, transcriptional regulators such as regulators of the LysR family and histone-like proteins (H-NS) also modulate gene expression 112. One such TCS, QseBC, was previously inferred to mediate resistance to a tachyplesin analogue by upregulating efflux genes based on transcriptomic analysis and hypersusceptibility of ΔqseBΔqseC mutants 113. Accordingly, we demonstrated that ΔqseBC mutants display only high accumulators of tachyplesin-NBD, suggesting that this TCS might play an important role in the control of tachyplesin efflux, possibly by regulating the expression of AcrA and TolC or other efflux components 64. This provides experimental support for the involvement of QseBC-mediated efflux of tachyplesin extending previous transcription-based inferences. However, it is also conceivable that the deletion of QseBC has pleiotropic effects on other cellular mechanisms involved in tachyplesin accumulation. We also acknowledge that our data do not take into account post-transcriptional modifications that represent a second control point to survive external stressors.

Earlier version: Drugs that inhibit the efflux activity of bacterial pathogens have been shown to be a promising strategy for repurposing antimicrobials 113–115. An FDA-approved drug, sertraline, was found to be effective in preventing the formation of the low accumulator phenotype by inhibiting tachyplesin efflux and enhancing the eradication of stationary phase bacteria. These data strongly suggest that EPIs are a promising approach for developing new combination therapies against stationary phase bacteria, contrary to previous evidence suggesting their ineffectiveness against stationary phase bacteria 33. However, it should be noted that the concentration of sertraline in the plasma of patients using sertraline as an antidepressant is below the concentration that we employed in our study 116. This limitation underscores the broader challenge of identifying EPIs that are both effective and minimally toxic within clinically achievable concentrations, while also meeting key therapeutic criteria such as broad-spectrum efficacy against diverse efflux pumps, high specificity for bacterial targets, and non-inducers of AMR 117. However, advances in biochemical, computational, and structural methodologies hold the potential to guide rational drug design, making the search for effective EPIs more promising 118. Therefore, more investigation should be carried out to further optimise the use of sertraline or other EPIs in combination with tachyplesin and other AMPs.

Revised version: Drugs that inhibit the efflux activity of bacterial pathogens have been shown to be a promising strategy for repurposing antimicrobials 114–116. An FDA-approved drug, sertraline, was found to be effective in preventing the formation of the low accumulator phenotype by inhibiting tachyplesin efflux and enhancing the eradication of stationary phase bacteria, although it is also conceivable that sertraline prevented the formation of the low accumulator phenotype via efflux independent mechanisms. These data strongly suggest that EPIs are a promising approach for developing new combination therapies against stationary phase bacteria, contrary to previous evidence suggesting their ineffectiveness against stationary phase bacteria 33. However, it should be noted that the concentration of sertraline in the plasma of patients using sertraline as an antidepressant is below the concentration that we employed in our study 117. This limitation underscores the broader challenge of identifying EPIs that are both effective and minimally toxic within clinically achievable concentrations, while also meeting key therapeutic criteria such as broad-spectrum efficacy against diverse efflux pumps, high specificity for bacterial targets, and non-inducers of AMR 118. However, advances in biochemical, computational, and structural methodologies hold the potential to guide rational drug design, making the search for effective EPIs more promising 119. Therefore, more investigation should be carried out to further optimise the use of sertraline or other EPIs in combination with tachyplesin and other AMPs.

Earlier version: Moreover, growth rate and metabolism have been shown to influence antimicrobial efficacy 119, with starved or slow growing stationary phase bacteria becoming transiently insensitive to antimicrobials when metabolites or ATP become unavailable 120. For example, expression of efflux genes has been shown to be increased in E. coli in minimal media compared to rich media 121. Due to the heightened resistance of starved and stationary phase cells, successful attempts have been made to selectively target metabolically dormant bacteria 122 or to sensitise recalcitrant cells via metabolites as a potential therapeutic strategy 123–125. Likewise, we found that stationary phase cells supplemented with glucose and casamino acids were more sensitive to treatment with tachyplesin. However, nutrient supplementation may potentially lead to an increase in bacterial cell count and disease burden, as well as an escalation in selection pressure on surviving cells. Moreover, nutrient supplementation as a therapeutic strategy may not be viable in many infection contexts, as host density and the immune system often regulate access to nutrients 3. It may be worthwhile to test other metabolites such as fumarate, which has received FDA approval in several drugs and has been shown to potentiate antimicrobial efficacy 124. Therefore, we conclude that the use of EPIs in combination with tachyplesin represents a more viable antimicrobial treatment against antimicrobial-refractory stationary phase bacteria.

Revised version: Moreover, growth rate and metabolism have been shown to influence antimicrobial efficacy 120, with starved or slow growing stationary phase bacteria becoming transiently insensitive to antimicrobials when metabolites or ATP become unavailable 121. For example, expression of efflux genes has been shown to be increased in E. coli in minimal media compared to rich media 122. Due to the heightened resistance of starved and stationary phase cells, successful attempts have been made to selectively target metabolically dormant bacteria 123 or to sensitise recalcitrant cells via metabolites as a potential therapeutic strategy 124–126. Likewise, we found that stationary phase cells supplemented with glucose and casamino acids were more sensitive to treatment with tachyplesin. However, nutrient supplementation may potentially lead to an increase in bacterial cell count and disease burden, as well as an escalation in selection pressure on surviving cells. Moreover, nutrient supplementation as a therapeutic strategy may not be viable in many infection contexts, as host density and the immune system often regulate access to nutrients 3. It may be worthwhile to test other metabolites such as fumarate, which has received FDA approval in several drugs and has been shown to potentiate antimicrobial efficacy 125. Therefore, we conclude that the use of EPIs in combination with tachyplesin represents a more viable antimicrobial treatment against antimicrobial-refractory stationary phase bacteria.

Methods

Chemicals and cell culture

All chemicals were purchased from Fisher Scientific or Sigma-Aldrich unless otherwise stated. LB medium (10 g L-1 tryptone, 5 g L-1 yeast extract and 10 g L-1 NaCl) and LB agar plates (LB with 15 g L- 1 agar) were used for planktonic growth and streak plates, respectively. Tachyplesin-1 and tachyplesin-1-NBD were synthesised by WuXi AppTech (Shanghai, China). Stock solutions of antimicrobials (1.28 mg mL-1), PI (1.5 mM), proteinase K (20 mg mL-1), and EtBr (1 mg mL-1) were obtained by dissolving in Milli-Q water. All EPIs and C12-resazurin were obtained by dissolving in dimethyl sulfoxide at a concentration of 1 mg mL-1. Carbon-free M9-minimal medium used for dilution of antimicrobials, PI, proteinase K, EtBr, EPIs, C12-resazurin, and bacteria was prepared using 5×M9 minimal salts (Merck, Germany), with an additional 2 mM MgSO4 and 0.1 mM CaCl2 in Milli-Q water. 0.4% glucose and 0.2% casamino acids were added to yield nutrient-supplemented M9. M9 was then filtered through a 0.22 μm Minisart® Syringe Filter (Sartorius, Germany).

Earlier version: E. coli BW25113 and S. aureus ATCC 25923 were purchased from Dharmacon™ (Horizon Discovery, UK). E. coli and K. pneumoniae (JS1187) clinical isolates were kindly provided by Fernanda Paganelli, UMC Utrecht, the Netherlands. P. aeruginosa (PA14 flgK::Tn5B30(Tc)) was kindly provided by George O’Toole, Dartmouth College, USA. E. coli single-gene deletion KO mutants were obtained from the Keio collection 126. To construct double gene-deletion mutants, the antibiotic resistance gene flanked by flippase recognition target sequences was removed from the chromosome using the pCP20 plasmid which encodes the flippase enzyme. The resulting transformants were then treated as described previously 127. Double gene-deletion mutants were then subsequently constructed by P1vir transduction of alleles from the Keio collection, as previously detailed 128. All strains were stored in a 50% glycerol stock at -80 °C. Streak plates for each strain were produced by thawing a small aliquot of the corresponding glycerol stock every week and plated onto LB agar plates. Stationary phase cultures were prepared by inoculating 100 mL fresh LB medium with a single bacterial colony from a streak plate and incubated on a shaking platform at 200 rpm and 37 °C for 17 h 129. Exponential phase cultures were prepared by transferring 100 μL of a stationary phase culture into 100 mL LB and incubated on a shaking platform at 200 rpm and 37 °C for 3 h.

Revised version: E. coli BW25113 and S. aureus ATCC 25923 were purchased from Dharmacon™ (Horizon Discovery, UK). E. coli and K. pneumoniae (JS1187) clinical isolates were kindly provided by Fernanda Paganelli, UMC Utrecht, the Netherlands. P. aeruginosa (PA14 flgK::Tn5B30(Tc)) was kindly provided by George O’Toole, Dartmouth College, USA. E. coli single-gene deletion KO mutants were obtained from the Keio collection 127. To construct double gene-deletion mutants, the antibiotic resistance gene flanked by flippase recognition target sequences was removed from the chromosome using the pCP20 plasmid which encodes the flippase enzyme. The resulting transformants were then treated as described previously 128. Double gene-deletion mutants were then subsequently constructed by P1vir transduction of alleles from the Keio collection, as previously detailed 129. All strains were stored in a 50% glycerol stock at -80 °C. Streak plates for each strain were produced by thawing a small aliquot of the corresponding glycerol stock every week and plated onto LB agar plates. Stationary phase cultures were prepared by inoculating 100 mL fresh LB medium with a single bacterial colony from a streak plate and incubated on a shaking platform at 200 rpm and 37 °C for 17 h 130. Exponential phase cultures were prepared by transferring 100 μL of a stationary phase culture into 100 mL LB and incubated on a shaking platform at 200 rpm and 37 °C for 3 h.

Synthesis of fluorescent antimicrobial derivatives

Earlier version: Fluorescent antimicrobial derivatives of tachyplesin 36,40, arenicin 130, polymyxin-B 131, and octapeptin132 were designed and synthesised based on structure-activity-relationship studies and synthetic protocols reported in previous publications 133, substituting a non-critical amino acid residue with an azidolysine residue that was then employed for the subsequent Cu-catalysed azide-alkyne cycloaddition ‘click’ reactions with nitrobenzoxadiazole (NBD)-alkyne. Detailed synthesis and characterisation of these probes will be reported in due course.

Revised version: Fluorescent antimicrobial derivatives of tachyplesin 36,40, arenicin 131, polymyxin-B 132, and octapeptin 133 were designed and synthesised based on structure-activity-relationship studies and synthetic protocols reported in previous publications 134, substituting a non-critical amino acid residue with an azidolysine residue that was then employed for the subsequent Cu-catalysed azide-alkyne cycloaddition ‘click’ reactions with nitrobenzoxadiazole (NBD)-alkyne. Detailed synthesis and characterisation of these probes will be reported in due course.

Phylogenetic and AMR genes analyses

Earlier version: The E. coli BW25113 genome sequence was downloaded from the National Center for Biotechnology Information’s (NCBI) GenBank® 134. E. coli clinical isolates MA02514, MA07534, MA08141, MC04960, and MC42862 were collected and sequenced by UMC Utrecht, the Netherlands 135. Briefly, genomic DNA libraries were prepared using the Nextera XT Library Prep Kit (Illumina, USA) and sequenced on either the Illumina MiSeq or NextSeq (Illumina, USA). Contigs were assembled using SPAdes Genome Assembler version 3.6.2 136, and contigs larger than 500 bp with at least 10× coverage were further analysed. E. coli clinical isolates JS1060, JS1073, JS1147, and JS1147 were collected by UMC Utrecht, the Netherlands, and were sequenced by a microbial genome sequencing service (MicrobesNG, UK). Briefly, genomic DNA libraries were prepared using the Nextera XT Library Prep Kit (Illumina, USA) following the manufacturer’s protocol with the following modifications: input DNA was increased two-fold, and PCR elongation time is increased to 45 s. DNA quantification and library preparation were carried out on a Hamilton Microlab STAR automated liquid handling system (Hamilton Bonaduz AG, Switzerland). Libraries were sequenced on an Illumina NovaSeq 6000 (Illumina, USA) using a 250 bp paired end protocol. Reads and adapter trimmed using Trimmomatic version 0.30 137 with a sliding window quality cutoff of Q15. De novo assembly was performed on samples using SPAdes version 3.7136, and contigs were annotated using Prokka 1.11 138.

Revised version: The E. coli BW25113 genome sequence was downloaded from the National Center for Biotechnology Information’s (NCBI) GenBank® 135. E. coli clinical isolates MA02514, MA07534, MA08141, MC04960, and MC42862 were collected and sequenced by UMC Utrecht, the Netherlands 136. Briefly, genomic DNA libraries were prepared using the Nextera XT Library Prep Kit (Illumina, USA) and sequenced on either the Illumina MiSeq or NextSeq (Illumina, USA). Contigs were assembled using SPAdes Genome Assembler version 3.6.2 137, and contigs larger than 500_bp with at least 10× coverage were further analysed. E. coli clinical isolates JS1060, JS1073, JS1147, and JS1147 were collected by UMC Utrecht, the Netherlands, and were sequenced by a microbial genome sequencing service (MicrobesNG, UK). Briefly, genomic DNA libraries were prepared using the Nextera XT Library Prep Kit (Illumina, USA) following the manufacturer’s protocol with the following modifications: input DNA was increased two-fold, and PCR elongation time is increased to 45 s. DNA quantification and library preparation were carried out on a Hamilton Microlab STAR automated liquid handling system (Hamilton Bonaduz AG, Switzerland). Libraries were sequenced on an Illumina NovaSeq 6000 (Illumina, USA) using a 250 bp paired end protocol. Reads and adapter trimmed using Trimmomatic version 0.30 138 with a sliding window quality cutoff of Q15. De novo assembly was performed on samples using SPAdes version 3.7 137, and contigs were annotated using Prokka 1.11 139.

Earlier version: Genomic data for AMR genes were screened through ABRicate 139 and the NCBI’s AMRFinderPlus databases 140. A rooted phylogenetic tree was created from the genomic data by pairwise distance estimation of all genomes and an outgroup genome (NCBI’s reference genome for Escherichia fergusonii) with the dist function from the Mash package 141, with a k-mer size of 21. Unrooted phylogenetic trees were created from the resulting pairwise distance matrices using the nj (neighbour joining) function in APE 142. These trees were then rooted to the outgroup (ape::root). The phylogenetic tree was plotted against presence/absence heatmaps of AMR genes using ggtree 143.

Revised version: Genomic data for AMR genes were screened through ABRicate 140 and the NCBI’s AMRFinderPlus databases 141. A rooted phylogenetic tree was created from the genomic data by pairwise distance estimation of all genomes and an outgroup genome (NCBI’s reference genome for Escherichia fergusonii) with the dist function from the Mash package 142, with a k-mer size of 21. Unrooted phylogenetic trees were created from the resulting pairwise distance matrices using the nj (neighbour joining) function in APE 143. These trees were then rooted to the outgroup (ape::root). The phylogenetic tree was plotted against presence/absence heatmaps of AMR genes using ggtree 144.

Flow cytometric assays

Bacterial cultures were prepared as described above. For all flow cytometric assays, bacterial cultures (stationary and exponential phase) were adjusted to an OD600 of 4 for E. coli and K. pneumoniae, 5 for S. aureus, and 2 for P. aeruginosa. Cultures were pelleted by centrifugation in a SLS4600 microcentrifuge (Scientific Laboratory Supplies, UK) at 4,000 ×g and room temperature for 5 min. The supernatant was then removed, and cells were resuspended in the appropriate volume of the treatment solution for incubation, with a final volume between 50 and 550 μL depending on the experiment. All treatments were incubated at 37 °C and 1,000 rpm on a ThermoMixer® C (Eppendorf, Germany), except for cold treatments, which were incubated on ice, for the stated duration. All wash steps were performed to remove any unaccumulated compounds by centrifugation at 4000 ×g for 5 min, after which the supernatant was removed to eliminate any extracellular drug.

For AMP accumulation assays, bacteria were incubated with the stated concentration of peptide in either carbon-free or nutrient-supplemented M9. For proteinase K degradation and C12-resazurin metabolic activity assays, E. coli was first incubated in 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 60 min followed by a wash step. Cells were then incubated in either 20 μg mL-1 (0.7 μM) proteinase K, or 1 μM C12-resazurin and 50 μM CCCP, all diluted in carbon-free M9. For tachyplesin-NBD efflux assays, E. coli was incubated in 46 μg mL-1 (18.2 μM) tachyplesin-NBD for 5 min, followed by a wash step and incubation in EPI diluted in carbon-free M9. For EtBr efflux assays, E. coli was preloaded with EtBr by adding 100 μg mL-1 (254 μM) EtBr to the stationary phase culture 90 min before a total of 17 h incubation (i.e. 15.5 h post-inoculation). Cells were then washed and incubated in 46 μg mL-1 (20.3 μM) unlabelled tachyplesin-1 diluted in carbon-free M9. Carbon-free M9 was used as the untreated control in all flow cytometric assays.

At the stated incubation time points, a 30 μL sacrificial aliquot was transferred to microcentrifuge tubes. A wash step was immediately performed, after which cells were resuspended and diluted 500× in carbon-free M9 before flow cytometry measurements. For the accumulation assays, to account for the time between sample collection and flow cytometry measurements, 10 min was added to the reported incubation time for the first two time points (i.e. the sacrificial aliquot taken immediately after tachyplesin-NBD addition was reported as 10 min).

Flow cytometric measurements were performed on the CytoFLEX S Flow Cytometer (Beckman Coulter, USA) equipped with a 488 nm (50 mW) and a 405 nm (80 mW) laser. Fluorescence of individual bacteria was measured using the fluorescein isothiocyanate (FITC) channel (488 nm 525/40BP) for fluorescent peptides and phycoerythrin (PE) channel (488 nm 585/42BP) for EtBr and C12-resorufin. Cell size was measured using FSC (forward scatter) and Violet SSC (side scatter) channels. Avalanche photodiode gains of FSC: 1,000, SSC: 500, Violet SSC: 1, FITC: 250 and a threshold value of SSC-A: 10,000 to limit background noise was used. Bacteria were gated to separate cells from background noise by plotting FSC-A and Violet SSC-A (Figure S1A). Background noise was then further separated based on cellular autofluorescence measured on the FITC-A channel for cells not treated with fluorescent-peptides (Figure S1B). An additional gate on the FITC-A channel was then used for cells treated with fluorescent-peptides to further separate background noise (Figure S1C). CytoFLEX Sheath Fluid (Beckman Coulter, USA) was used as sheath fluid. Data was collected using CytExpert software (Beckman Coulter, USA) and exported to FlowJo™ version 10.9 software (BD Biosciences, USA) for analysis.

Microfluidic mother machine device fabrication

Earlier version: Microfluidic mother machine devices were fabricated by pouring a 10:1 (base:curing agent) polydimethylsiloxane (PDMS) mixture 144 into an epoxy mould kindly provided by S. Jun 49. Each mother machine device contains approximately 6,000 bacterial hosting channels with a width, height, and length of 1, 1.5, and 25 µm, respectively. These channels are connected to a main microfluidic chamber with a width and height of 25 and 100 µm, respectively. After degassing the PDMS mixture in a vacuum desiccator (SP Bel-Art, USA) with a two-stage rotary vane pump FRVP series (Fisher Scientific, USA), the PDMS was cured at 70 °C for 2 h and peeled from the epoxy mould to obtain 12 individual chips145. Fluidic inlet and outlet were achieved using a 0.75 mm biopsy punch (WellTech Labs, Taiwan) at the two ends of the main chamber of the mother machine. The PDMS chip was then washed with ethanol and dried with nitrogen gas, followed by removal of small particles using adhesive tape (3M, USA) along with a rectangular glass coverslip (Fisher Scientific, USA). The microfluidic device was assembled by exposing the PDMS chip and glass coverslip to air plasma treatment at 30 W plasma power for 10 s with the Zepto One plasma cleaner (Diener Electronic, Germany) and then placed in contact to irreversibly bind the PDMS chip and glass coverslip 146. The mother machine was then filled with 5 μL of 50 mg mL -1 (752 μM) bovine serum albumin and incubated at 37 °C for 1 h to passivate the charge within the channels after plasma treatment, screening electrostatic interaction between bacterial membranes, PDMS and glass surfaces 147.

Revised version: Microfluidic mother machine devices were fabricated by pouring a 10:1 (base:curing agent) polydimethylsiloxane (PDMS) mixture 145 into an epoxy mould kindly provided by S. Jun 49. Each mother machine device contains approximately 6,000 bacterial hosting channels with a width, height, and length of 1, 1.5, and 25 µm, respectively. These channels are connected to a main microfluidic chamber with a width and height of 25 and 100 µm, respectively. After degassing the PDMS mixture in a vacuum desiccator (SP Bel-Art, USA) with a two-stage rotary vane pump FRVP series (Fisher Scientific, USA), the PDMS was cured at 70 °C for 2 h and peeled from the epoxy mould to obtain 12 individual chips 146. Fluidic inlet and outlet were achieved using a 0.75 mm biopsy punch (WellTech Labs, Taiwan) at the two ends of the main chamber of the mother machine. The PDMS chip was then washed with ethanol and dried with nitrogen gas, followed by removal of small particles using adhesive tape (3M, USA) along with a rectangular glass coverslip (Fisher Scientific, USA). The microfluidic device was assembled by exposing the PDMS chip and glass coverslip to air plasma treatment at 30 W plasma power for 10 s with the Zepto One plasma cleaner (Diener Electronic, Germany) and then placed in contact to irreversibly bind the PDMS chip and glass coverslip 147. The mother machine was then filled with 5 μL of 50 mg mL -1 (752 μM) bovine serum albumin and incubated at 37 °C for 1 h to passivate the charge within the channels after plasma treatment, screening electrostatic interaction between bacterial membranes, PDMS and glass surfaces 148.

Microfluidics-microscopy assay to measure tachyplesin-NBD accumulation and PI staining in individual E. coli

Earlier version: A stationary phase E. coli culture was prepared as described above. The culture was adjusted to OD600 75 by centrifuging 50 mL of culture in a conical tube at 3,220 ×g at room temperature for 5 min in a 5810 R centrifuge (Eppendorf, Germany). The supernatant was removed, and pellet resuspended in carbon-free M9 and vortexed. 5 μL of bacterial suspension was then injected into the mother machine device using a pipette and incubated at 37 °C for 30 to 60 min to allow for filling of ∼80% of bacteria hosting channels 148. The loaded microfluidic device was then mounted on an Olympus IX73 inverted microscope (Olympus, Japan) connected to a 60×, 1.2 N.A. objective UPLSAPO60XW/1.20 (Olympus, Japan) and a Zyla 4.2 sCMOS camera (Andor, UK) for visual inspection. Each region of interest was adjusted to visualise 23 bacteria hosting channels and a mean of 11 regions of interest were selected for imaging for each experiment via automated stages M-545.USC and P-545.3C7 (Physik Instrumente, Germany) for coarse and fine movements, respectively. Fluorinated ethylene propylene tubing (1/32” × 0.0008”) was then inserted into both inlet and outlet accesses as previously reported 149. The inlet tubing was connected to a Flow Unit S flow rate sensor (Fluigent, France) and MFCS-4C pressure control system (Fluigent, France) controlled via MAESFLO 3.3.1 software (Fluigent, France) allowing for computerised, accurate regulation of fluid flow into the microfluidic device. A 300 μL h-1 flow of carbon-free M9 for 8 min was used to clear the main channel of excess bacteria that had not entered the bacterial hosting channels. Tachyplesin-NBD (46 μg mL-1 or 18.2 μM) with or without sertraline (30 μg mL-1 or 98 μM) in carbon-free M9 was then flowed through the microfluidic device at 300 μL h-1 for 8 min, before reducing to 100 μL h-1 until 240 min. Carbon-free M9 was then flushed through the device at 300 μL h-1 for 60 min to wash away unaccumulated tachyplesin-NBD. PI staining was performed by flowing PI (1.5 mM) at 300 μL h-1 for 15 min. Brightfield and fluorescence images were captured consecutively every 10 min until the end of the experiment by switching between brightfield and fluorescence mode in a custom LabVIEW software (National Instruments, USA) module. Fluorescence images were captured by illuminating bacteria with a pE-300white broad-spectrum LED (CoolLED, UK) for 0.06 s on the blue excitation band at 20% intensity with a FITC filter for tachyplesin-NBD and green excitation band at 100 % intensity with a DAPI/TEXAS filter for PI. The entire experiment was performed at 37 °C in an environmental chamber (Solent Scientific, UK) enclosing the microscope and microfluidics equipment. A step-by-step guide to the process above can be found in 150.

Revised version: A stationary phase E. coli culture was prepared as described above. The culture was adjusted to OD600 75 by centrifuging 50 mL of culture in a conical tube at 3,220 ×g at room temperature for 5 min in a 5810 R centrifuge (Eppendorf, Germany). The supernatant was removed, and pellet resuspended in carbon-free M9 and vortexed. 5 μL of bacterial suspension was then injected into the mother machine device using a pipette and incubated at 37 °C for 30 to 60 min to allow for filling of ∼80% of bacteria hosting channels 149. The loaded microfluidic device was then mounted on an Olympus IX73 inverted microscope (Olympus, Japan) connected to a 60×, 1.2 N.A. objective UPLSAPO60XW/1.20 (Olympus, Japan) and a Zyla 4.2 sCMOS camera (Andor, UK) for visual inspection. Each region of interest was adjusted to visualise 23 bacteria hosting channels and a mean of 11 regions of interest were selected for imaging for each experiment via automated stages M-545.USC and P-545.3C7 (Physik Instrumente, Germany) for coarse and fine movements, respectively. Fluorinated ethylene propylene tubing (1/32” × 0.0008”) was then inserted into both inlet and outlet accesses as previously reported 150. The inlet tubing was connected to a Flow Unit S flow rate sensor (Fluigent, France) and MFCS-4C pressure control system (Fluigent, France) controlled via MAESFLO 3.3.1 software (Fluigent, France) allowing for computerised, accurate regulation of fluid flow into the microfluidic device. A 300 μL h-1 flow of carbon-free M9 for 8 min was used to clear the main channel of excess bacteria that had not entered the bacterial hosting channels. Tachyplesin-NBD (46 μg mL-1 or 18.2 μM) with or without sertraline (30 μg mL-1 or 98 μM) in carbon-free M9 was then flowed through the microfluidic device at 300 μL h-1 for 8 min, before reducing to 100 μL h-1 until 240 min. Carbon-free M9 was then flushed through the device at 300 μL h-1 for 60 min to wash away unaccumulated tachyplesin-NBD. PI staining was performed by flowing PI (1.5 mM) at 300 μL h-1 for 15 min. Brightfield and fluorescence images were captured consecutively every 10 min until the end of the experiment by switching between brightfield and fluorescence mode in a custom LabVIEW software (National Instruments, USA) module. Fluorescence images were captured by illuminating bacteria with a pE-300white broad-spectrum LED (CoolLED, UK) for 0.06 s on the blue excitation band at 20% intensity with a FITC filter for tachyplesin-NBD and green excitation band at 100 % intensity with a DAPI/TEXAS filter for PI. The entire experiment was performed at 37 °C in an environmental chamber (Solent Scientific, UK) enclosing the microscope and microfluidics equipment. A step-by-step guide to the process above can be found in 151.

Image and data analyses

Earlier version: Images were processed using ImageJ software as previously described 151. Each individual bacterium was tracked throughout the 240 min tachyplesin-NBD treatment, 60 min carbon-free M9 wash and 15 min PI staining. To measure tachyplesin-NBD and PI accumulation, a rectangle was drawn around each bacterium in each brightfield image at every time point, obtaining its length, width, and relative position in the microfluidic hosting channel. The same rectangle was then overlaid onto the corresponding fluorescence image to measure the mean fluorescence intensity, which represents the total fluorescence normalised by cell size (i.e., the area covered by each bacterium in the 2D images), thus accounting for variations in tachyplesin-NBD and PI accumulation due to the cell cycle 152. The same rectangle was then shifted to the nearest adjacent channel that did not contain any bacteria to measure the mean background fluorescence from extracellular tachyplesin-NBD in the media. This mean background fluorescence value was subsequently subtracted from the bacterium’s mean fluorescence value. To measure tachyplesin-NBD fluorescence profiles, the fluorescence values of a straight horizontal line spanning 20 pixels (left to right) were recorded then converted and presented in μm. The membrane over cell centre fluorescence ratios were calculated by dividing the fluorescence of the membrane by the fluorescence value of the centre of the cell. The mean membrane fluorescence was obtained using the fluorescence of the seventh and thirteenth pixel, and the cell centre fluorescence using the tenth pixel along the profile measurements. To measure cell sizes of low and high accumulators, the vertical length of the rectangles drawn to measure fluorescence was taken for each cell then converted and presented in μm. All data were then analysed and plotted in Prism version 10.2.0 (GraphPad Software, USA).

Revised version: Images were processed using ImageJ software as previously described 152. Each individual bacterium was tracked throughout the 240 min tachyplesin-NBD treatment, 60 min carbon-free M9 wash and 15 min PI staining. To measure tachyplesin-NBD and PI accumulation, a rectangle was drawn around each bacterium in each brightfield image at every time point, obtaining its length, width, and relative position in the microfluidic hosting channel. The same rectangle was then overlaid onto the corresponding fluorescence image to measure the mean fluorescence intensity, which represents the total fluorescence normalised by cell size (i.e., the area covered by each bacterium in the 2D images), thus accounting for variations in tachyplesin-NBD and PI accumulation due to the cell cycle 153. The same rectangle was then shifted to the nearest adjacent channel that did not contain any bacteria to measure the mean background fluorescence from extracellular tachyplesin-NBD in the media. This mean background fluorescence value was subsequently subtracted from the bacterium’s mean fluorescence value. To measure tachyplesin-NBD fluorescence profiles, the fluorescence values of a straight horizontal line spanning 20 pixels (left to right) were recorded then converted and presented in μm. The membrane over cell centre fluorescence ratios were calculated by dividing the fluorescence of the membrane by the fluorescence value of the centre of the cell. The mean membrane fluorescence was obtained using the fluorescence of the seventh and thirteenth pixel, and the cell centre fluorescence using the tenth pixel along the profile measurements. To measure cell sizes of low and high accumulators, the vertical length of the rectangles drawn to measure fluorescence was taken for each cell then converted and presented in μm. All data were then analysed and plotted in Prism version 10.2.0 (GraphPad Software, USA).

Statistical significance tests, including the ordinary one-way ANOVA with Tukey’s multiple comparisons tests, paired two-tailed Wilcoxon nonparametric test, unpaired two-tailed nonparametric Mann-Whitney U test, as well as linear and nonlinear (semi-log) regressions were calculated using Prism version 10.2.0 (GraphPad Software, USA).

Fluorescence-activated cell sorting

Samples for fluorescence-activated cell sorting were generated following the same protocol as flow cytometry accumulation assays. Cells were incubated for 60 min and a wash step carried out before performing a 30× dilution on 500 μL of treated cells for cell sorting. Cell sorting was performed on a BD FACSAria™ Fusion Flow Cytometer, equipped with a 488 nm (50 mW) laser (BD Biosciences, USA). Drops were generated at 70 psi sheath pressure, 87.0 kHz frequency, and 5.6 amplitude using a 70 μm nozzle. BD FACSFlow™ (BD Biosciences, USA) was used as sheath fluid and the instrument was set up for sorting with the BD™ Cytometer Setup & Tracking Beads Kit, and BD FACS™ Accudrop Beads reagents (BD Biosciences, USA), following the manufacturer’s instructions. Cells were sorted based on their fluorescence measured on the 488 nm 530/30BP channel with photomultiplier tubes (PMT) voltages of FSC: 452, SSC: 302 and 488 nm 530/30BP: 444 and a threshold value of 200 was applied on the SSC channel to limit background noise. Cell aggregates were excluded by plotting SSC-W against SSC-H (Figure S8A). Tachyplesin-NBD treated cells were sorted into two groups (low and high accumulators) and the untreated control treatment through a control gate (Figure S8E and S8C, respectively). A one-drop sorting envelope with purity rules was used. Approximately 1,000,000 million cells for each group were sorted directly into RNAprotect® (Qiagen, Germany) to stabilise RNA in preparation for extraction and into carbon-free M9 for lipid Folch extractions. An additional combined sample containing a 1:1, v/v mixture of sorted low and high accumulators was generated (referred to as 1:1 mix accumulators) to aid and validate downstream transcriptomics analyses. Cells were sorted into carbon-free M9 for cell viability measurements following the same protocol above. Data was collected via BD FACSDiva™ version 8.0.1 (BD Biosciences, USA) software and analysed using FlowJo™ version 10.9 software (BD Biosciences, USA). Post-sort analyses revealed high levels of purity and very low occurrences of events measured outside of the sorting gate for sorted cells (Figure S8D and S8F-H). Increased events measured in the M9 gate in post-sort samples are due to a reduced concentration of cells in post-sort analyses.

Cell viability assays

Cell viability was assessed by diluting cells in carbon-free M9, then plating on LB agar plates. Plates were then incubated at 37 °C for 17 h followed by colony forming unit counts. For flow cytometry accumulation and time-kill assays, sacrificial 30 µL aliquots of cells were obtained from the treatment microcentrifuge tube at their respective time points for dilution and spread plating. For cells sorted via fluorescence-activated cell sorting, low and high tachyplesin-NBD accumulators, and untreated control cells were sorted into carbon-free M9 before dilution and spread plating.

RNA extraction and sequencing

Earlier version: RNA extractions were performed on cells sorted into RNAprotect® (Qiagen, Germany). Enzymatic lysis and proteinase K digestion of bacteria was performed following protocol 4 in the RNAprotect® Bacteria Reagent Handbook (Qiagen, Germany). Purification of total RNA from bacterial lysate using the RNeasy® Mini Kit (Qiagen, Germany) following protocol 7 in the handbook. Due to the small quantity of initial material from cell sorting, RNeasy MinElute® spin columns were used instead of RNeasy® Mini spin columns. A further on-column DNase digestion using the RNase-Free DNase set (Qiagen, Germany) was performed following appendix B in the handbook. RNA concentration and quality were assessed using the Agilent High Sensitivity RNA ScreenTape System (Agilent, USA) following the provided protocol. Samples returned a mean RNA concentration of 2.5 ng μL-1 and mean RNA integrity number equivalent (RINe) score of 7.1. rRNA depletion was performed with the Illumina® Stranded Total RNA Prep, Ligation with Ribo-Zero Plus kit (Illumina, USA) following the manufacturer’s instructions and sequencing carried out on the Illumina NovaSeq™ 6000 (Illumina, USA). Transcript abundance was quantified using Salmon for each gene in all samples. Differential gene expression was performed with DESeq2 in R software to quantify the log2 fold-change in transcript reads 153 for each gene and subsequent principal component analysis using DESeq2 and a built-in R method (prcomp) 52.

Revised version: RNA extractions were performed on cells sorted into RNAprotect® (Qiagen, Germany). Enzymatic lysis and proteinase K digestion of bacteria was performed following protocol 4 in the RNAprotect® Bacteria Reagent Handbook (Qiagen, Germany). Purification of total RNA from bacterial lysate using the RNeasy® Mini Kit (Qiagen, Germany) following protocol 7 in the handbook. Due to the small quantity of initial material from cell sorting, RNeasy MinElute® spin columns were used instead of RNeasy® Mini spin columns. A further on-column DNase digestion using the RNase-Free DNase set (Qiagen, Germany) was performed following appendix B in the handbook. RNA concentration and quality were assessed using the Agilent High Sensitivity RNA ScreenTape System (Agilent, USA) following the provided protocol. Samples returned a mean RNA concentration of 2.5 ng μL-1 and mean RNA integrity number equivalent (RINe) score of 7.1. rRNA depletion was performed with the Illumina® Stranded Total RNA Prep, Ligation with Ribo-Zero Plus kit (Illumina, USA) following the manufacturer’s instructions and sequencing carried out on the Illumina NovaSeq™ 6000 (Illumina, USA). Transcript abundance was quantified using Salmon for each gene in all samples. Differential gene expression was performed with DESeq2 in R software to quantify the log2 fold-change in transcript reads 154 for each gene and subsequent principal component analysis using DESeq2 and a built-in R method (prcomp) 52.

Cluster and gene ontology analyses

Earlier version: For cluster and gene ontology analysis, differential expression was tested with edgeR (version 3.28.1)154. Predicted log fold-changes were positively correlated between replicates, however a batch effect was detected, and replicate number was retained as a model term. Transcripts with low expression were filtered out of the data via edgeR before fitting the differential expression models. Clustering analysis was performed using the mclust package (version 5.4.7) for R 155. Only transcripts with differential expression FDR < 0.05 in at least one cell type were subjected to clustering analysis 156. All variance structures were tested for a range of 2 to 20 clusters and the minimal, best-fitting model was identified by the Bayes information criterion. Gene ontology enrichment analysis was performed using the clusterProfiler package (version 4.10.0) for R 157,158. Enrichment in terms belonging to the “Biological Process” ontology was calculated for each gene cluster, relative to the set of all genes quantified in the experiment, via a one-sided Fisher exact test (hypergeometric test). P-values were adjusted for false discovery by using the method of Benjamini and Hochberg. Finally, the lists of significantly enriched terms were simplified to remove redundant terms, as assessed via their semantic similarity to other enriched terms, using clusterProfiler’s simplify function.

Revised version: For cluster and gene ontology analysis, differential expression was tested with edgeR (version 3.28.1) 155. Predicted log fold-changes were positively correlated between replicates, however a batch effect was detected, and replicate number was retained as a model term. Transcripts with low expression were filtered out of the data via edgeR before fitting the differential expression models. Clustering analysis was performed using the mclust package (version 5.4.7) for R 156. Only transcripts with differential expression FDR < 0.05 in at least one cell type were subjected to clustering analysis 157. All variance structures were tested for a range of 2 to 20 clusters and the minimal, best-fitting model was identified by the Bayes information criterion. Gene ontology enrichment analysis was performed using the clusterProfiler package (version 4.10.0) for R 158,159. Enrichment in terms belonging to the “Biological Process” ontology was calculated for each gene cluster, relative to the set of all genes quantified in the experiment, via a one-sided Fisher exact test (hypergeometric test). P-values were adjusted for false discovery by using the method of Benjamini and Hochberg. Finally, the lists of significantly enriched terms were simplified to remove redundant terms, as assessed via their semantic similarity to other enriched terms, using clusterProfiler’s simplify function.

Transcription factor activity

Earlier version: Relative activities of transcription factors were estimated from differential expression results (generated in edgeR) using VIPER (version 1.2.0) for R 61. Transcription factor regulons were derived from the full RegulonDB database (version 10.9) 159, fixing the likelihood to 1. Only transcripts with differential expression FDR < 0.05 in at least one cell type were used to estimate transcription factor activities. VIPER was run with a minimum allowed regulon size of 20 (minsize=20), on the full gene data set (eset.filter=FALSE) and was set to normalise enrichment scores (nes=TRUE).

Revised version: Relative activities of transcription factors were estimated from differential expression results (generated in edgeR) using VIPER (version 1.2.0) for R 61. Transcription factor regulons were derived from the full RegulonDB database (version 10.9) 160, fixing the likelihood to 1. Only transcripts with differential expression FDR < 0.05 in at least one cell type were used to estimate transcription factor activities. VIPER was run with a minimum allowed regulon size of 20 (minsize=20), on the full gene data set (eset.filter=FALSE) and was set to normalise enrichment scores (nes=TRUE).

Lipid extraction and lipidomic analysis

Earlier version: Lipid extractions were performed following a modified version of the Folch extraction 56,160. Briefly, 100 μL of sample was added to 250 μL of methanol and 125 μL of chloroform in a microcentrifuge tube. Samples were incubated for 60 min and vortexed every 15 min. Then, 380 μL of chloroform and 90 μL of 0.2 M potassium chloride was added to each sample. Cells were centrifuged at 14,000 ×g for 10 min to obtain a lipophilic phase which was transferred to a glass vial and dried under a nitrogen stream.

Revised version: Lipid extractions were performed following a modified version of the Folch extraction 56,161. Briefly, 100 μL of sample was added to 250 μL of methanol and 125 μL of chloroform in a microcentrifuge tube. Samples were incubated for 60 min and vortexed every 15 min. Then, 380 μL of chloroform and 90 μL of 0.2 M potassium chloride was added to each sample. Cells were centrifuged at 14,000 ×g for 10 min to obtain a lipophilic phase which was transferred to a glass vial and dried under a nitrogen stream.

The lipophilic phase was then resuspended in 20 μL of a methanol:chloroform solution (1:1 v/v), then 980 μL of an isopropanol:acetonitrile:water solution (2:1:1, v/v/v) was added. Samples were analysed on the Agilent 6560 Q-TOF-MS (Agilent, USA) coupled with the Agilent 1290 Infinity II LC system (Agilent, USA). An aliquot of 0.5 μL for positive ionisation mode and 2 μL for negative ionisation mode from each sample was injected into a Kinetex® 5 μm EVO C18 100 A, 150 mm × 2.1 μm column (Phenomenex, USA). The column was maintained at 50 °C at a flow rate of 0.3 mL min-1. For the positive ionisation mode, the mobile phases consisted of (A) acetonitrile:water (2:3 v/v) with ammonium formate (10 mM) and B acetonitrile:isopropanol (1:9, v/v) with ammonium formate (10 mM). For the negative ionisation mode, the mobile phases consisted of (A) acetonitrile:water (2:3 v/v) with ammonium acetate (10 mM) and (B) acetonitrile:isopropanol (1:9, v/v) with ammonium acetate (10 mM). The chromatographic separation was obtained with the following gradient: 0-1 min 70% B; 1-3.5 min 86% B; 3.5-10 min 86% B; 10.1-17 min 100% B; 17.1-10 min 70% B. The mass spectrometry platform was equipped with an Agilent Jet Stream Technology Ion Source (Agilent, USA), which was operated in both positive and negative ion modes with the following parameters: gas temperature (200 °C); gas flow (nitrogen), 10 L min-1; nebuliser gas (nitrogen), 50 psig; sheath gas temperature (300 °C); sheath gas flow, 12 L min-1; capillary voltage 3500 V for positive, and 3000 V for negative; nozzle voltage 0 V; fragmentor 150 V; skimmer 65 V, octupole RF 7550 V; mass range, 40-1,700 m/z; capillary voltage, 3.5 kV; collision energy 20 eV in positive, and 25 eV in negative mode. MassHunter software (Agilent, USA) was used for instrument control.

Low and high tachyplesin-NBD accumulators, and untreated control cells sorted via cell sorting were analysed using quadrupole time-of-flight liquid chromatography/mass spectrometry (Q-TOF-LC/MS). Mass Profiler 10.0 (Agilent, USA) was used to generate a matrix containing lipid features across all samples, then filtered by standard deviation, and normalised by sum calculations. Multivariate statistical analyses were performed on these features using SIMCA® software 15.0 (Sartorius, Germany). First, a principal component analysis (PCA) was performed followed by a partial least square-discriminant analysis (PLS-DA) with its orthogonal extension (OPLS-DA), which was used to visualise differences between low and high tachyplesin-NBD accumulators and untreated control samples.

Earlier version: Tachyplesin-NBD treated and untreated cells were analysed using ultra-high performance liquid chromatography-quadrupole time-of-flight mass spectrometry (UHPLC-Q-TOF-MS) and detected a mean of 1,000 mass spectrometry features in positive and negative ionisation modes. Lipid annotation was performed following COSMOS consortium guidelines 161 with the CEU Mass Mediator version 3.0 online tool 162. Lipid Annotator software (Agilent, USA) was then used to identify the lipid classes and their relative abundances (Figure 3E). Lipids with variable importance in projection (VIP) scores above 1 were annotated (Table S2).

Revised version: Tachyplesin-NBD treated and untreated cells were analysed using ultra-high performance liquid chromatography-quadrupole time-of-flight mass spectrometry (UHPLC-Q-TOF-MS) and detected a mean of 1,000 mass spectrometry features in positive and negative ionisation modes. Lipid annotation was performed following COSMOS consortium guidelines 162 with the CEU Mass Mediator version 3.0 online tool 163. Lipid Annotator software (Agilent, USA) was then used to identify the lipid classes and their relative abundances (Figure 3E). Lipids with variable importance in projection (VIP) scores above 1 were annotated (Table S2).

Determination of minimum inhibitory concentration

Earlier version: Exponential phase bacterial cultures were prepared as described above and grown to an OD600 of 0.5. 100 μL LB broth was added to all wells of a 96-well microtiter plate (Sarstedt, Germany). Then, 100 μL of a 2048 μg mL-1 (6693.5 μM) sertraline solution was added to the first nine wells of the top row and mixed. Sertraline was then serially diluted to 8 μg mL-1 (26.1 μM) by transferring 100 μL to the next row, mixing thoroughly after each transfer, using a multi-channel pipette. After the final dilution step, 100 μL was discarded to maintain equal volumes across wells. Exponential phase bacterial cultures were then diluted to yield 106 colony-forming units (CFU) mL-1 and added to all wells containing sertraline. Each plate included nine positive control wells (i.e. bacteria in LB without sertraline), and fifteen negative control wells (i.e. LB only, without bacteria or sertraline). Plates were covered with a lid and incubated at 37 °C for 24 h. The MIC of sertraline against E. coli BW25113 was determined using the CLARIOstar microplate reader (BMG Labtech, Germany) by measuring OD600 after incubation. The MIC was defined as the lowest concentration of sertraline at which no visible growth was observed, as determined by comparison to positive control wells.

Revised version: Exponential phase bacterial cultures were prepared as described above and grown to an OD600 of 0.5. 100 μL LB broth was added to all wells of a 96-well microtiter plate (Sarstedt, Germany). Then, 100-1 μL of a 2048 μg mL (6693.5 μM) sertraline solution was added to the first nine wells of the top row and mixed. Sertraline was then serially diluted to 8 μg mL-1 (26.1 μM) by transferring 100 μL to the next row, mixing thoroughly after each transfer, using a multi-channel pipette. After the final dilution step, 100 μL was discarded to maintain equal volumes across wells. Exponential phase bacterial cultures were then diluted to yield 106 colony-forming units (CFU) mL-1 and added to all wells containing sertraline. Each plate included nine positive control wells (i.e. bacteria in LB without sertraline), and fifteen negative control wells (i.e. LB only, without bacteria or sertraline). Plates were covered with a lid and incubated at 37 °C for 24 h. The MIC of sertraline against E. coli BW25113 was determined using the CLARIOstar microplate reader (BMG Labtech, Germany) by measuring OD600 after incubation. The MIC was defined as the lowest concentration of sertraline at which no visible growth was observed, as determined by comparison to positive control wells.

Accession Numbers

Sequencing data associated with this study are linked under NCBI BioProject accession number PRJNA1096674 (https://www.ncbi.nlm.nih.gov/bioproject/PRJNA1096674).